A Myelin Map of Trunk Folds in the Elephant Trigeminal Nucleus

Author response:

The following is the authors’ response to the original reviews.

We are thankful for the handling of our manuscript. The following is a summary of our response and what we have done:

(1) We are most thankful for the very thorough evaluation of our manuscript.

(2) We were a bit shocked by the very negative commentary of referee 2.

(3) We think, what put referee 2 off so much is that we were overconfident in the strength of our conclusions. We consider such overconfidence a big mistake. We have revised the manuscript to fix this problem.

(4) We respond in great depth to all criticism and also go into technicalities.

(5) We consider the possibility of a mistake. Yet, we carefully weighed the evidence advanced by referee 2 and by us and found that a systematic review supports our conclusions. Hence, we also resist the various attempts to crush our paper.

(6) We added evidence (peripherin-antibody staining; our novel Figure 2) that suggests we correctly identified the inferior olive.

(7) The eLife format – in which critical commentary is published along with the paper – is a fantastic venue to publish, what appears to be a surprisingly controversial issue.

eLife assessment

This potentially valuable study uses classic neuroanatomical techniques and synchrotron X-ray tomography to investigate the mapping of the trunk within the brainstem nuclei of the elephant brain. Given its unique specializations, understanding the somatosensory projections from the elephant trunk would be of general interest to evolutionary neurobiologists, comparative neuroscientists, and animal behavior scientists. However, the anatomical analysis is inadequate to support the authors' conclusion that they have identified the elephant trigeminal sensory nuclei rather than a different brain region, specifically the inferior olive.

Comment: We are happy that our paper is considered to be potentially valuable. Also, the editors highlight the potential interest of our work for evolutionary neurobiologists, comparative neuroscientists, and animal behavior scientists. The editors are more negative when it comes to our evidence on the identification of the trigeminal nucleus vs the inferior olive. We have five comments on this assessment. (i) We think this assessment is heavily biased by the comments of referee 2. We show that the referee’s comments are more about us than about our paper. Hence, the referee failed to do their job (refereeing our paper) and should not have succeeded in leveling our paper. (ii) We have no ad hoc knock-out experiments to distinguish the trigeminal nucleus vs the inferior olive. Such experiments (extracellular recording & electrolytic lesions, viral tracing would be done in a week in mice, but they cannot and should not be done in elephants. (iii) We have extraordinary evidence. Nobody has ever described a similarly astonishing match of body (trunk folds) and myeloarchitecture in the brain before. (iv) We show that our assignment of the trigeminal nucleus vs the inferior olive is more plausible than the current hypothesis about the assignment of the trigeminal nucleus vs the inferior olive as defended by referee 2. We think this is why it is important to publish our paper. (v) We think eLife is the perfect place for our publication because the deviating views of referee 2 are published along.

Change: We performed additional peripherin-antibody staining to differentiate the inferior olive and trigeminal nucleus. Peripherin is a cytoskeletal protein that is found in peripheral nerves and climbing fibers. Specifically, climbing fibers of various species (mouse, rabbit, pig, cow, and human; Errante et al., 1998) are stained intensely with peripherin-antibodies. What is tricky for our purposes is that there is also some peripherin-antibody reactivity in the trigeminal nuclei (Errante et al., 1998). Such peripherin-antibody reactivity is weaker, however, and lacks the distinct axonal bundle signature that stems from the strong climbing fiber peripherin-reactivity as seen in the inferior olive (Errante et al., 1998). As can be seen in our novel Figure 2, we observe peripherin-reactivity in axonal bundles (i.e. in putative climbing fibers), in what we think is the inferior olive. We also observe weak peripherin-reactivity, in what we think is the trigeminal nucleus, but not the distinct and strong labeling of axonal bundles. These observations are in line with our ideas but are difficult to reconcile with the views of the referee. Specifically, the lack of peripherin-reactive axon bundles suggests that there are no climbing fibers in what the referee thinks is the inferior olive.

Errante, L., Tang, D., Gardon, M., Sekerkova, G., Mugnaini, E., & Shaw, G. (1998). The intermediate filament protein peripherin is a marker for cerebellar climbing fibres. Journal of neurocytology, 27, 69-84.

Reviewer #1 :

Summary:

This fundamental study provides compelling neuroanatomical evidence underscoring the sensory function of the trunk in African and Asian elephants. Whereas myelinated tracts are classically appreciated as mediating neuronal connections, the authors speculate that myelinated bundles provide functional separation of trunk folds and display elaboration related to the "finger" projections. The authors avail themselves of many classical neuroanatomical techniques (including cytochrome oxidase stains, Golgi stains, and myelin stains) along with modern synchrotron X-ray tomography. This work will be of interest to evolutionary neurobiologists, comparative neuroscientists, and the general public, with its fascinating exploration of the brainstem of an icon sensory specialist.

Comment: We are incredibly grateful for this positive assessment.

Changes: None.

Strengths:

- The authors made excellent use of the precious sample materials from 9 captive elephants.

- The authors adopt a battery of neuroanatomical techniques to comprehensively characterize the structure of the trigeminal subnuclei and properly re-examine the "inferior olive".

- Based on their exceptional histological preparation, the authors reveal broadly segregated patterns of metabolic activity, similar to the classical "barrel" organization related to rodent whiskers.

Comment: The referee provides a concise summary of our findings.

Changes: None.

Weaknesses:

- As the authors acknowledge, somewhat limited functional description can be provided using histological analysis (compared to more invasive techniques).

- The correlation between myelinated stripes and trunk fold patterns is intriguing, and Figure 4 presents this idea beautifully. I wonder - is the number of stripes consistent with the number of trunk folds? Does this hold for both species?

Comment: We agree with the referee’s assessment. We note that cytochrome-oxidase staining is an at least partially functional stain, as it reveals constitutive metabolic activity. A significant problem of the work in elephants is that our recording possibilities are limited, which in turn limits functional analysis. As indicated in Figure 5 (our former Figure 4) for the African elephant Indra, there was an excellent match of trunk folds and myelin stripes. Asian elephants have more, and less conspicuous trunk folds than African elephants. As illustrated in Figure 7, Asian elephants have more, and less conspicuous myelin stripes. Thus, species differences in myelin stripes correlate with species differences in trunk folds.

Changes: We clarify the relation of myelin stripe and trunk fold patterns in our description of Figure 7.

Reviewer #2 (Public Review):

The authors describe what they assert to be a very unusual trigeminal nuclear complex in the brainstem of elephants, and based on this, follow with many speculations about how the trigeminal nuclear complex, as identified by them, might be organized in terms of the sensory capacity of the elephant trunk.

Comment: We agree with the referee’s assessment that the putative trigeminal nucleus described in our paper is highly unusual in size, position, vascularization, and myeloarchitecture. This is why we wrote this paper. We think these unusual features reflect the unique facial specializations of elephants, i.e. their highly derived trunk. Because we have no access to recordings from the elephant brainstem, we cannot back up all our functional interpretations with electrophysiological evidence; it is therefore fair to call them speculative.

Changes: None.

The identification of the trigeminal nuclear complex/inferior olivary nuclear complex in the elephant brainstem is the central pillar of this manuscript from which everything else follows, and if this is incorrect, then the entire manuscript fails, and all the associated speculations become completely unsupported.

Comment: We agree.

Changes: None.

The authors note that what they identify as the trigeminal nuclear complex has been identified as the inferior olivary nuclear complex by other authors, citing Shoshani et al. (2006; 10.1016/j.brainresbull.2006.03.016) and Maseko et al (2013; 10.1159/000352004), but fail to cite either Verhaart and Kramer (1958; PMID 13841799) or Verhaart (1962; 10.1515/9783112519882-001). These four studies are in agreement, but the current study differs.

Comment & Change: We were not aware of the papers of Verhaart and included them in the revised manusript.

Let's assume for the moment that the four previous studies are all incorrect and the current study is correct. This would mean that the entire architecture and organization of the elephant brainstem is significantly rearranged in comparison to ALL other mammals, including humans, previously studied (e.g. Kappers et al. 1965, The Comparative Anatomy of the Nervous System of Vertebrates, Including Man, Volume 1 pp. 668-695) and the closely related manatee (10.1002/ar.20573). This rearrangement necessitates that the trigeminal nuclei would have had to "migrate" and shorten rostrocaudally, specifically and only, from the lateral aspect of the brainstem where these nuclei extend from the pons through to the cervical spinal cord (e.g. the Paxinos and Watson rat brain atlases), the to the spatially restricted ventromedial region of specifically and only the rostral medulla oblongata. According to the current paper, the inferior olivary complex of the elephant is very small and located lateral to their trigeminal nuclear complex, and the region from where the trigeminal nuclei are located by others appears to be just "lateral nuclei" with no suggestion of what might be there instead.

Comment: We have three comments here:

(1) The referee correctly notes that we argue the elephant brainstem underwent fairly major rearrangements. In particular, we argue that the elephant inferior olive was displaced laterally, by a very large cell mass, which we argue is an unusually large trigeminal nucleus. To our knowledge, such a large compact cell mass is not seen in the ventral brain stem of any other mammal.

(2) The referee makes it sound as if it is our private idea that the elephant brainstem underwent major rearrangements and that the rest of the evidence points to a conventional ‘rodent-like’ architecture. This is far from the truth, however. Already from the outside appearance (see our Figure 1B and Figure 7A) it is clear that the elephant brainstem has huge ventral bumps not seen in any other mammal. An extraordinary architecture also holds at the organizational level of nuclei. Specifically, the facial nucleus – the most carefully investigated nucleus in the elephant brainstem – has an appearance distinct from that of the facial nuclei of all other mammals (Maseko et al., 2013; Kaufmann et al., 2022). If both the overall shape and the constituting nuclei of the brainstem are very different from other mammals, it is very unlikely if not impossible that the elephant brainstem follows in all regards a conventional ‘rodent-like’ architecture.

(3) The inferior olive is an impressive nucleus in the partitioning scheme we propose (Figure 2). In fact – together with the putative trigeminal nucleus we describe – it’s the most distinctive nucleus in the elephant brainstem. We have not done volumetric measurements and cell counts here, but think this is an important direction for future work. What has informed our work is that the inferior olive nucleus we describe has the serrated organization seen in the inferior olive of all mammals. We will discuss these matters in depth below.

Changes: None.

Such an extraordinary rearrangement of brainstem nuclei would require a major transformation in the manner in which the mutations, patterning, and expression of genes and associated molecules during development occur. Such a major change is likely to lead to lethal phenotypes, making such a transformation extremely unlikely. Variations in mammalian brainstem anatomy are most commonly associated with quantitative changes rather than qualitative changes (10.1016/B978-0-12-804042-3.00045-2).

Comment: We have two comments here:

(1) The referee claims that it is impossible that the elephant brainstem differs from a conventional brainstem architecture because this would lead to lethal phenotypes etc. Following our previous response, this argument does not hold. It is out of the question that the elephant brainstem looks very different from the brainstem of other mammals. Yet, it is also evident that elephants live. The debate we need to have is not if the elephant brainstem differs from other mammals, but how it differs from other mammals.

(2) In principle we agree with the referee’s thinking that the model of the elephant brainstem that is most likely to be correct is the one that requires the least amount of rearrangements to other mammals. We therefore prepared a comparison of the model the referee is proposing (Maseko et al., 2013; see Referee Table 1 below) with our proposition. We scored these models on their similarity to other mammals. We find that the referee’s ideas (Maseko et al., 2013) require more rearrangements relative to other mammals than our suggestion.

Changes: Inclusion of Referee Table 1, which we discuss in depth below.

The impetus for the identification of the unusual brainstem trigeminal nuclei in the current study rests upon a previous study from the same laboratory (10.1016/j.cub.2021.12.051) that estimated that the number of axons contained in the infraorbital branch of the trigeminal nerve that innervate the sensory surfaces of the trunk is approximately 400 000. Is this number unusual? In a much smaller mammal with a highly specialized trigeminal system, the platypus, the number of axons innervating the sensory surface of the platypus bill skin comes to 1 344 000 (10.1159. Yet, there is no complex rearrangement of the brainstem trigeminal nuclei in the brain of the developing or adult platypus (Ashwell, 2013, Neurobiology of Monotremes), despite the brainstem trigeminal nuclei being very large in the platypus (10.1159/000067195). Even in other large-brained mammals, such as large whales that do not have a trunk, the number of axons in the trigeminal nerve ranges between 400,000 and 500,000 (10.1007. The lack of comparative support for the argument forwarded in the previous and current study from this laboratory, and that the comparative data indicates that the brainstem nuclei do not change in the manner suggested in the elephant, argues against the identification of the trigeminal nuclei as outlined in the current study. Moreover, the comparative studies undermine the prior claim of the authors, informing the current study, that "the elephant trigeminal ganglion ... point to a high degree of tactile specialization in elephants" (10.1016/j.cub.2021.12.051). While clearly, the elephant has tactile sensitivity in the trunk, it is questionable as to whether what has been observed in elephants is indeed "truly extraordinary".

Comment: These comments made us think that the referee is not talking about the paper we submitted, but that the referee is talking about us and our work in general. Specifically, the referee refers to the platypus and other animals dismissing our earlier work, which argued for a high degree of tactile specialization in elephants. We think the referee’s intuitions are wrong and our earlier work is valid.

Changes: We prepared a Author response image 1 (below) that puts the platypus brain, a monkey brain, and the elephant trigeminal ganglion (which contains a large part of the trunk innervating cells) in perspective.

The elephant trigeminal ganglion is comparatively large.

Platypus brain, monkey brain, and elephant ganglion. The elephant has two trigeminal ganglia, which contain the first-order somatosensory neurons. They serve mainly for tactile processing and are large compared to a platypus brain (from the comparative brain collection) and are similar in size to a monkey brain. The idea that elephants might be highly specialized for trunk touch is also supported by the analysis of the sensory nerves of these animals (Purkart et al., 2022). Specifically, we find that the infraorbital nerve (which innervates the trunk) is much thicker than the optic nerve (which mediates vision) and the vestibulocochlear nerve (which mediates hearing). Thus, not everything is large about elephants; instead, the data argue that these animals are heavily specialized for trunk touch.

But let's look more specifically at the justification outlined in the current study to support their identification of the unusually located trigeminal sensory nuclei of the brainstem.

(1) Intense cytochrome oxidase reactivity.

(2) Large size of the putative trunk module.

(3) Elongation of the putative trunk module.

(4) The arrangement of these putative modules corresponds to elephant head

anatomy.

(5) Myelin stripes within the putative trunk module that apparently match trunk folds.

(6) Location apparently matches other mammals.

(7) Repetitive modular organization apparently similar to other mammals.

(8) The inferior olive described by other authors lacks the lamellated appearance of this structure in other mammals.

Comment: We agree those are key issues.

Changes: None.

Let's examine these justifications more closely.

(1) Cytochrome oxidase histochemistry is typically used as an indicative marker of neuronal energy metabolism. The authors indicate, based on the "truly extraordinary" somatosensory capacities of the elephant trunk, that any nuclei processing this tactile information should be highly metabolically active, and thus should react intensely when stained for cytochrome oxidase. We are told in the methods section that the protocols used are described by Purkart et al (2022) and Kaufmann et al (2022). In neither of these cited papers is there any description, nor mention, of the cytochrome oxidase histochemistry methodology, thus we have no idea of how this histochemical staining was done. To obtain the best results for cytochrome oxidase histochemistry, the tissue is either processed very rapidly after buffer perfusion to remove blood or in recently perfusion-fixed tissue (e.g., 10.1016/0165-0270(93)90122-8). Given: (1) the presumably long post-mortem interval between death and fixation - "it often takes days to dissect elephants"; (2) subsequent fixation of the brains in 4% paraformaldehyde for "several weeks"; (3) The intense cytochrome oxidase reactivity in the inferior olivary complex of the laboratory rat (Gonzalez-Lima, 1998, Cytochrome oxidase in neuronal metabolism and Alzheimer's diseases); and (4) The lack of any comparative images from other stained portions of the elephant brainstem; it is difficult to support the justification as forwarded by the authors. The histochemical staining observed is likely background reactivity from the use of diaminobenzidine in the staining protocol. Thus, this first justification is unsupported.

Comment: The referee correctly notes the description of our cytochrome-oxidase reactivity staining was lacking. This is a serious mistake of ours for which we apologize very much. The referee then makes it sound as if we messed up our cytochrome-oxidase staining, which is not the case. All successful (n = 3; please see our technical comments in the recommendation section) cytochrome-oxidase stainings were done with elephants with short post-mortem times (≤ 2 days) to brain removal/cooling and only brief immersion fixation (≤ 1 day). Cytochrome-oxidase reactivity in elephant brains appears to be more sensitive to quenching by fixation than is the case for rodent brains. We think it is a good idea to include a cytochrome-oxidase staining overview picture because we understood from the referee’s comments that we need to compare our partitioning scheme of the brainstem with that of other authors. To this end, we add a cytochrome-oxidase staining overview picture (Author response image 3) along with an alternative interpretation from Maseko et al., 2013.

Changes: (1) We added details on our cytochrome-oxidase reactivity staining protocol and the cytochrome-oxidase reactivity in the elephant brain in the manuscript and in our response to the general recommendations.

(2) We provide a detailed discussion of the technicalities of cytochrome-oxidase staining below in the recommendation section, where the referee raised further criticisms.

(3) We include a cytochrome-oxidase staining overview picture (Author response image 2) along with an alternative interpretation from Maseko et al., 2013.

Cytochrome-oxidase staining overview.

Coronal cytochrome-oxidase staining overview from African elephant cow Indra; the section is taken a few millimeters posterior to the facial nucleus. Brown is putatively neural cytochrome-reactivity, and white is the background. Black is myelin diffraction and (seen at higher resolution, when you zoom in) erythrocyte cytochrome-reactivity in blood vessels (see our Figure 1E-G); such blood vessel cytochrome-reactivity is seen, because we could not perfuse the animal. There appears to be a minimal outside-in-fixation artifact (i.e. a more whitish/non-brownish appearance of the section toward the borders of the brain). This artifact is not seen in sections from Indra that we processed earlier or in other elephant brains processed at shorter post-mortem/fixation delays (see our Figure 1C).

The same structures can be recognized in Author response image 2 and Supplememntary figure 36 of Maseko et al. (2013). The section is taken at an anterior-posterior level, where we encounter the trigeminal nuclei in pretty much all mammals. Note that the neural cytochrome reactivity is very high, in what we refer to as the trigeminal-nuclei-trunk-module and what Maseko et al. refer to as inferior olive. Myelin stripes can be recognized here as white omissions.

At the same time, the cytochrome-oxidase-reactivity is very low in what Maseko et al. refer to as trigeminal nuclei. The indistinct appearance and low cytochrome-oxidase-reactivity of the trigeminal nuclei in the scheme of Maseko et al. (2013) is unexpected because trigeminal nuclei stain intensely for cytochrome-oxidase-reactivity in most mammals and because the trigeminal nuclei represent the elephant’s most important body part, the trunk. Staining patterns of the trigeminal nuclei as identified by Maseko et al. (2013) are very different at more posterior levels; we will discuss this matter below.

Justifications (2), (3), and (4) are sequelae from justification (1). In this sense, they do not count as justifications, but rather unsupported extensions.

Comment: These are key points of our paper that the referee does not discuss.

Changes: None.

(4) and (5) These are interesting justifications, as the paper has clear internal contradictions, and (5) is a sequelae of (4). The reader is led to the concept that the myelin tracts divide the nuclei into sub-modules that match the folding of the skin on the elephant trunk. One would then readily presume that these myelin tracts are in the incoming sensory axons from the trigeminal nerve. However, the authors note that this is not the case: "Our observations on trunk module myelin stripes are at odds with this view of myelin. Specifically, myelin stripes show no tapering (which we would expect if axons divert off into the tissue). More than that, there is no correlation between myelin stripe thickness (which presumably correlates with axon numbers) and trigeminal module neuron numbers. Thus, there are numerous myelinated axons, where we observe few or no trigeminal neurons. These observations are incompatible with the idea that myelin stripes form an axonal 'supply' system or that their prime function is to connect neurons. What do myelin stripe axons do, if they do not connect neurons? We suggest that myelin stripes serve to separate rather than connect neurons." So, we are left with the observation that the myelin stripes do not pass afferent trigeminal sensory information from the "truly extraordinary" trunk skin somatic sensory system, and rather function as units that separate neurons - but to what end? It appears that the myelin stripes are more likely to be efferent axonal bundles leaving the nuclei (to form the olivocerebellar tract). This justification is unsupported.

Comment: The referee cites some of our observations on myelin stripes, which we find unusual. We stand by the observations and comments. The referee does not discuss the most crucial finding we report on myelin stripes, namely that they correspond remarkably well to trunk folds.

Changes: None.

(6) The authors indicate that the location of these nuclei matches that of the trigeminal nuclei in other mammals. This is not supported in any way. In ALL other mammals in which the trigeminal nuclei of the brainstem have been reported they are found in the lateral aspect of the brainstem, bordered laterally by the spinal trigeminal tract. This is most readily seen and accessible in the Paxinos and Watson rat brain atlases. The authors indicate that the trigeminal nuclei are medial to the facial nerve nucleus, but in every other species, the trigeminal sensory nuclei are found lateral to the facial nerve nucleus. This is most salient when examining a close relative, the manatee (10.1002/ar.20573), where the location of the inferior olive and the trigeminal nuclei matches that described by Maseko et al (2013) for the African elephant. This justification is not supported.

Comment: The referee notes that we incorrectly state that the position of the trigeminal nuclei matches that of other mammals. We think this criticism is justified.

Changes: We prepared a comparison of the Maseko et al. (2013) scheme of the elephant brainstem with our scheme of the elephant brainstem (see below Referee Table 1). Here we acknowledge the referee’s argument and we also changed the manuscript accordingly.

(7) The dual to quadruple repetition of rostrocaudal modules within the putative trigeminal nucleus as identified by the authors relies on the fact that in the neurotypical mammal, there are several trigeminal sensory nuclei arranged in a column running from the pons to the cervical spinal cord, these include (nomenclature from Paxinos and Watson in roughly rostral to caudal order) the Pr5VL, Pr5DM, Sp5O, Sp5I, and Sp5C. However, these nuclei are all located far from the midline and lateral to the facial nerve nucleus, unlike what the authors describe in the elephants. These rostrocaudal modules are expanded upon in Figure 2, and it is apparent from what is shown that the authors are attributing other brainstem nuclei to the putative trigeminal nuclei to confirm their conclusion. For example, what they identify as the inferior olive in Figure 2D is likely the lateral reticular nucleus as identified by Maseko et al (2013). This justification is not supported.

Comment: The referee again compares our findings to the scheme of Maseko et al. (2013) and rejects our conclusions on those grounds. We think such a comparison of our scheme is needed, indeed.

Changes: We prepared a comparison of the Maseko et al. (2013) scheme of the elephant brainstem with our scheme of the elephant brainstem (see below Referee Table 1).

(8) In primates and related species, there is a distinct banded appearance of the inferior olive, but what has been termed the inferior olive in the elephant by other authors does not have this appearance, rather, and specifically, the largest nuclear mass in the region (termed the principal nucleus of the inferior olive by Maseko et al, 2013, but Pr5, the principal trigeminal nucleus in the current paper) overshadows the partial banded appearance of the remaining nuclei in the region (but also drawn by the authors of the current paper). Thus, what is at debate here is whether the principal nucleus of the inferior olive can take on a nuclear shape rather than evince a banded appearance. The authors of this paper use this variance as justification that this cluster of nuclei could not possibly be the inferior olive. Such a "semi-nuclear/banded" arrangement of the inferior olive is seen in, for example, giraffe (10.1016/j.jchemneu.2007.05.003), domestic dog, polar bear, and most specifically the manatee (a close relative of the elephant) (brainmuseum.org; 10.1002/ar.20573). This justification is not supported.

Comment: We carefully looked at the brain sections referred to by the referee in the brainmuseum.org collection. We found contrary to the referee’s claims that dogs, polar bears, and manatees have a perfectly serrated (a cellular arrangement in curved bands) appearance of the inferior olive. Accordingly, we think the referee is not reporting the comparative evidence fairly and we wonder why this is the case.

Changes: None.

Thus, all the justifications forwarded by the authors are unsupported. Based on methodological concerns, prior comparative mammalian neuroanatomy, and prior studies in the elephant and closely related species, the authors fail to support their notion that what was previously termed the inferior olive in the elephant is actually the trigeminal sensory nuclei. Given this failure, the justifications provided above that are sequelae also fail. In this sense, the entire manuscript and all the sequelae are not supported.

Comment: We disagree. To summarize:

(1) Our description of the cytochrome oxidase staining lacked methodological detail, which we have now added; the cytochrome oxidase reactivity data are great and support our conclusions.

(2)–(5)The referee does not really discuss our evidence on these points.

(6) We were wrong and have now fixed this mistake.

(7) The referee asks for a comparison to the Maseko et al. (2013) scheme (agreed, see Referee Table 1).

(8) The referee bends the comparative evidence against us.

Changes: None.

A comparison of the elephant brainstem partitioning schemes put forward by Maseko et al 2013 and by Reveyaz et al.

To start with, we would like to express our admiration for the work of Maseko et al. (2013). These authors did pioneering work on obtaining high-quality histology samples from elephants. Moreover, they made a heroic neuroanatomical effort, in which they assigned 147 brain structures to putative anatomical entities. Most of their data appear to refer to staining in a single elephant and one coronal sectioning plane. The data quality and the illustration of results are excellent.

We studied mainly two large nuclei in six (now 7) elephants in three (coronal, parasagittal, and horizontal) sectioning planes. The two nuclei in question are the two most distinct nuclei in the elephant brainstem, namely an anterior ventromedial nucleus (the trigeminal trunk module in our terminology; the inferior olive in the terminology of Maseko et al., 2013) and a more posterior lateral nucleus (the inferior olive in our terminology; the posterior part of the trigeminal nuclei in the terminology of Maseko et al., 2013).

Author response image 3 gives an overview of the two partitioning schemes for inferior olive/trigeminal nuclei along with the rodent organization (see below).

Overview of the brainstem organization in rodents & elephants

The strength of the Maseko et al. (2013) scheme is the excellent match of the position of elephant nuclei to the position of nuclei in the rodent (Author response image 3). We think this positional match reflects the fact that Maseko et al. (2013) mapped a rodent partitioning scheme on the elephant brainstem. To us, this is a perfectly reasonable mapping approach. As the referee correctly points out, the positional similarity of both elephant inferior olive and trigeminal nuclei to the rodent strongly argues in favor of the Maseko et al. (2013), because brainstem nuclei are positionally very conservative.

Other features of the Maseko et al. (2013) scheme are less favorable. The scheme marries two cyto-architectonically very distinct divisions (an anterior indistinct part) and a super-distinct serrated posterior part to be the trigeminal nuclei. We think merging entirely distinct subdivisions into one nucleus is a byproduct of mapping a rodent partitioning scheme on the elephant brainstem. Neither of the two subdivisions resemble the trigeminal nuclei of other mammals. The cytochrome oxidase staining patterns differ markedly across the anterior indistinct part (see our Author response image 3) and the posterior part of the trigeminal nuclei and do not match with the intense cytochrome oxidase reactivity of other mammalian trigeminal nuclei (Author response image 2). Our anti-peripherin staining (the novel Figure 2 of our manuscript) indicates that there probably no climbing fibers, in what Maseko et al. think. is inferior olive; this is a potentially fatal problem for the hypothesis. The posterior part of Maseko et al. (2013) trigeminal nuclei has a distinct serrated appearance that is characteristic of the inferior olive in other mammals. Moreover, the inferior olive of Maseko et al. (2013) lacks the serrated appearance of the inferior olive seen in pretty much all mammals; this is a serious problem.

The partitioning scheme of Reveyaz et al. comes with poor positional similarity but avoids the other problems of the Maseko et al. (2013) scheme. Our explanation for the positionally deviating location of trigeminal nuclei is that the elephant grew one of the if not the largest trigeminal systems of all mammals. As a result, the trigeminal nuclei grew through the floor of the brainstem. We understand this is a post hoc just-so explanation, but at least it is an explanation.

The scheme of Reveyaz et al. was derived in an entirely different way from the Maseko model. Specifically, we were convinced that the elephant trigeminal nuclei ought to be very special because of the gigantic trigeminal ganglia (Purkart et al., 2022). Cytochrome-oxidase staining revealed a large distinct nucleus with an elongated shape. Initially, we were freaked out by the position of the nucleus and the fact that it was referred to as inferior olive by other authors. When we found an inferior-olive-like nucleus at a nearby (although at an admittedly unusual) location, we were less worried. We then optimized the visualization of myelin stripes (brightfield imaging etc.) and were able to collect an entire elephant trunk along with the brain (African elephant cow Indra). When we made the one-to-one match of Indra’s trunk folds and myelin stripes (former Figure 4, now Figure 5) we were certain that we had identified the trunk module of the trigeminal nuclei. We already noted at the outset of our rebuttal that we now consider such certainty a fallacy of overconfidence. In light of the comments of Referee 2, we feel that a further discussion of our ideas is warranted.

A strength of the Reveyaz model is that nuclei look like single anatomical entities. The trigeminal nuclei look like trigeminal nuclei of other mammals, the trunk module has a striking resemblance to the trunk and the inferior olive looks like the inferior olive of other mammals.

We evaluated the fit of the two models in the form of a table (Author response table 1; below). Unsurprisingly, Author response table 1 aligns with our views of elephant brainstem partitioning.

++ = Very attractive; + = attractive; - = unattractive; -- = very unattractive

We scored features that are clear and shared by all mammals – as far as we know them – as very attractive.

We scored features that are clear and are not shared by all mammals – as far as we know them – as very unattractive.

Attractive features are either less clear or less well-shared features.

Unattractive features are either less clear or less clearly not shared features.

Author response table 1 suggests two conclusions to us. (i) The Reveyaz et al. model has mainly favorable properties. The Maseko et al. (2013) model has mainly unfavorable properties. Hence, the Reveyaz et al. model is more likely to be true. (ii) The outcome is not black and white, i.e., both models have favorable and unfavorable properties. Accordingly, we overstated our case in our initial submission and toned down our claims in the revised manuscript.

What the authors have not done is to trace the pathway of the large trigeminal nerve in the elephant brainstem, as was done by Maseko et al (2013), which clearly shows the internal pathways of this nerve, from the branch that leads to the fifth mesencephalic nucleus adjacent to the periventricular grey matter, through to the spinal trigeminal tract that extends from the pons to the spinal cord in a manner very similar to all other mammals. Nor have they shown how the supposed trigeminal information reaches the putative trigeminal nuclei in the ventromedial rostral medulla oblongata. These are but two examples of many specific lines of evidence that would be required to support their conclusions. Clearly, tract tracing methods, such as cholera toxin tracing of peripheral nerves cannot be done in elephants, thus the neuroanatomy must be done properly and with attention to detail to support the major changes indicated by the authors.

Comment: The referee claims that Maseko et al. (2013) showed by ‘tract tracing’ that the structures they refer to trigeminal nuclei receive trigeminal input. This statement is at least slightly misleading. There is nothing of what amounts to proper ‘tract tracing’ in the Maseko et al. (2013) paper, i.e. tracing of tracts with post-mortem tracers. We tried proper post-mortem tracing but failed (no tracer transport) probably as a result of the limitations of our elephant material. What Maseko et al. (2013) actually did is look a bit for putative trigeminal fibers and where they might go. We also used this approach. In our hands, such ‘pseudo tract tracing’ works best in unstained material under bright field illumination, because myelin is very well visualized. In such material, we find: (i) massive fiber tracts descending dorsoventrally roughly from where both Maseko et al. 2013 and we think the trigeminal tract runs. (ii) These fiber tracts run dorsoventrally and approach, what we think is the trigeminal nuclei from lateral.

Changes: Ad hoc tract tracing see above.

So what are these "bumps" in the elephant brainstem?

Four previous authors indicate that these bumps are the inferior olivary nuclear complex. Can this be supported?

The inferior olivary nuclear complex acts "as a relay station between the spinal cord (n.b. trigeminal input does reach the spinal cord via the spinal trigeminal tract) and the cerebellum, integrating motor and sensory information to provide feedback and training to cerebellar neurons" (https://www.ncbi.nlm.nih.gov/books/NBK542242/). The inferior olivary nuclear complex is located dorsal and medial to the pyramidal tracts (which were not labeled in the current study by the authors but are clearly present in Fig. 1C and 2A) in the ventromedial aspect of the rostral medulla oblongata. This is precisely where previous authors have identified the inferior olivary nuclear complex and what the current authors assign to their putative trigeminal nuclei. The neurons of the inferior olivary nuclei project, via the olivocerebellar tract to the cerebellum to terminate in the climbing fibres of the cerebellar cortex.

Comment: We agree with the referee that in the Maseko et al. (2013) scheme the inferior olive is exactly where we expect it from pretty much all other mammals. Hence, this is a strong argument in favor of the Maseko et al. (2013) scheme and a strong argument against the partitioning scheme suggested by us.

Changes: Please see our discussion above.

Elephants have the largest (relative and absolute) cerebellum of all mammals (10.1002/ar.22425), this cerebellum contains 257 x109 neurons (10.3389/fnana.2014.00046; three times more than the entire human brain, 10.3389/neuro.09.031.2009). Each of these neurons appears to be more structurally complex than the homologous neurons in other mammals (10.1159/000345565; 10.1007/s00429-010-0288-3). In the African elephant, the neurons of the inferior olivary nuclear complex are described by Maseko et al (2013) as being both calbindin and calretinin immunoreactive. Climbing fibres in the cerebellar cortex of the African elephant are clearly calretinin immunopositive and also are likely to contain calbindin (10.1159/000345565). Given this, would it be surprising that the inferior olivary nuclear complex of the elephant is enlarged enough to create a very distinct bump in exactly the same place where these nuclei are identified in other mammals?

Comment: We agree with the referee that it is possible and even expected from other mammals that there is an enlargement of the inferior olive in elephants. Hence, a priori one might expect the ventral brain stem bumps to the inferior olive, this is perfectly reasonable and is what was done by previous authors. The referee also refers to calbindin and calretinin antibody reactivity. Such antibody reactivity is indeed in line with the referee’s ideas and we considered these findings in our Referee Table 1. The problem is, however, that neither calbindin nor calretinin antibody reactivity are highly specific and indeed both nuclei in discussion (trigeminal nuclei and inferior olive) show such reactivity. Unlike the peripherin-antibody staining advanced by us, calbindin nor calretinin antibody reactivity cannot distinguish the two hypotheses debated.

Changes: Please see our discussion above.

What about the myelin stripes? These are most likely to be the origin of the olivocerebellar tract and probably only have a coincidental relationship with the trunk. Thus, given what we know, the inferior olivary nuclear complex as described in other studies, and the putative trigeminal nuclear complex as described in the current study, is the elephant inferior olivary nuclear complex. It is not what the authors believe it to be, and they do not provide any evidence that discounts the previous studies. The authors are quite simply put, wrong. All the speculations that flow from this major neuroanatomical error are therefore science fiction rather than useful additions to the scientific literature.

Comment: It is unlikely that the myelin stripes are the origin of the olivocerebellar tract as suggested by the referee. Specifically, the lack of peripherin-reactivity indicates that these fibers are not climbing fibers (our novel Figure 2). In general, we feel the referee does not want to discuss the myelin stripes and obviously thinks we made up the strange correspondence of myelin stripes and trunk folds.

Changes: Please see our discussion above.

What do the authors actually have?

The authors have interesting data, based on their Golgi staining and analysis, of the inferior olivary nuclear complex in the elephant.

Comment: The referee reiterates their views.

Changes: None.

Reviewer #3 (Public Review):

Summary:

The study claims to investigate trunk representations in elephant trigeminal nuclei located in the brainstem. The researchers identified large protrusions visible from the ventral surface of the brainstem, which they examined using a range of histological methods. However, this ventral location is usually where the inferior olivary complex is found, which challenges the author's assertions about the nucleus under analysis. They find that this brainstem nucleus of elephants contains repeating modules, with a focus on the anterior and largest unit which they define as the putative nucleus principalis trunk module of the trigeminal. The nucleus exhibits low neuron density, with glia outnumbering neurons significantly. The study also utilizes synchrotron X-ray phase contrast tomography to suggest that myelin-stripe-axons traverse this module. The analysis maps myelin-rich stripes in several specimens and concludes that based on their number and patterning they likely correspond with trunk folds; however, this conclusion is not well supported if the nucleus has been misidentified.

Comment: The referee gives a concise summary of our findings. The referee acknowledges the depth of our analysis and also notes our cellular results. The referee – in line with the comments of Referee 2 – also points out that a misidentification of the nucleus under study is potentially fatal for our analysis. We thank the referee for this fair assessment.

Changes: We feel that we need to alert the reader more broadly to the misidentification concern. We think the critical comments of Referee 2, which will be published along with our manuscript, will go a long way in doing so. We think the eLife publishing format is fantastic in this regard. We will also include pointers to these concerns in the revised manuscript.

Strengths:

The strength of this research lies in its comprehensive use of various anatomical methods, including Nissl staining, myelin staining, Golgi staining, cytochrome oxidase labeling, and synchrotron X-ray phase contrast tomography. The inclusion of quantitative data on cell numbers and sizes, dendritic orientation and morphology, and blood vessel density across the nucleus adds a quantitative dimension. Furthermore, the research is commendable for its high-quality and abundant images and figures, effectively illustrating the anatomy under investigation.

Comment: Again, a very fair and balanced set of comments. We are thankful for these comments.

Changes: None.

Weaknesses:

While the research provides potentially valuable insights if revised to focus on the structure that appears to be the inferior olivary nucleus, there are certain additional weaknesses that warrant further consideration. First, the suggestion that myelin stripes solely serve to separate sensory or motor modules rather than functioning as an "axonal supply system" lacks substantial support due to the absence of information about the neuronal origins and the termination targets of the axons. Postmortem fixed brain tissue limits the ability to trace full axon projections. While the study acknowledges these limitations, it is important to exercise caution in drawing conclusions about the precise role of myelin stripes without a more comprehensive understanding of their neural connections.

Comment: The referee points out a significant weakness of our study, namely our limited understanding of the origin and targets of the axons constituting the myelin stripes. We are very much aware of this problem and this is also why we directed high-powered methodology like synchrotron X-ray tomograms to elucidate the structure of myelin stripes. Such analysis led to advances, i.e., we now think, what looks like stripes are bundles and we understand the constituting axons tend to transverse the module. Such advances are insufficient, however, to provide a clear picture of myelin stripe connectivity.

Changes: We think solving the problems raised by the referee will require long-term methodological advances and hence we will not be able to solve these problems in the current revision. Our long-term plans for confronting these issues are the following: (i) Improving our understanding of long-range connectivity by post-mortem tracing and MR-based techniques such as Diffusion-Tensor-Imaging. (ii) Improving our understanding of mid and short-range connectivity by applying even larger synchrotron X-ray tomograms and possible serial EM.

Second, the quantification presented in the study lacks comparison to other species or other relevant variables within the elephant specimens (i.e., whole brain or brainstem volume). The absence of comparative data for different species limits the ability to fully evaluate the significance of the findings. Comparative analyses could provide a broader context for understanding whether the observed features are unique to elephants or more common across species. This limitation in comparative data hinders a more comprehensive assessment of the implications of the research within the broader field of neuroanatomy. Furthermore, the quantitative comparisons between African and Asian elephant specimens should include some measure of overall brain size as a covariate in the analyses. Addressing these weaknesses would enable a richer interpretation of the study's findings.

Comment: The referee suggests another series of topics, which include the analysis of brain parts volumes or overall brain size. We agree these are important issues, but we also think such questions are beyond the scope of our study.

Changes: We hope to publish comparative data on elephant brain size and shape later this year.

Recommendations for the authors:

Reviewer #1 (Recommendations For The Authors):

I realize that elephant brains are a limiting resource in this project, along with the ability to perform functional investigations. However, I believe that Prof. Jon Kaas (Vanderbilt University) has one or more series of Nissl-stained brainstems from elephants. These might be of potential interest, as they were previously used to explore general patterns of trigeminal brainstem organization in a comparative manner (see Sawyer and Sarko, 2017, "Comparative Anatomy and Evolution of the Somatosensory Brain Stem" in the Evolution of Nervous System series) and might shed light on the positioning of the trigeminal complex and IO, with parts of the trigeminal nerve itself still attached to these sections.

Comment: The referee suggests adding data from more elephants and we think this is a great suggestion because our ns are small. We followed this advice. We agree we need more comparative neuroanatomy of elephants and the urgency of this matter is palpable in the heated debate we have with Referee 2. Specifically, we need more long-range and short-range analysis of elephant brains.

Changes: We plan to include data in the revised manuscript about cytoarchitectonics (Nissl), cytochrome-oxidase reactivity, and possibly also antibody reactivity from an additional animal, i.e., from the African elephant cow Bibi. The quality of this specimen is excellent and the post-mortem time to brain extraction was very short.

We also have further plans for connectivity analysis (see our response above), but such data will not become available fast enough for the revision.

Other recommendations:

- A general schematic showing input from trunk to PrV to the trigeminal subnuclei (as well as possibly ascending connections) might be informative to the reader, in terms of showing which neural relay is being examined.

Comment: We think this is a very good suggestion in principle, but we were not satisfied with the schematics we came up with.

Changes: None.

- Perhaps a few more sentences described the significance of synchrotron tomography for those who may be unfamiliar.

Comment & Change: We agree and implement this suggestion.

- "Belly-shaped" trunk module description is unclear on page 9.

Comment & Change: We clarified this matter.

- Typo on the last sentence of page 9.

Comment & Change: We fixed this mistake.

Reviewer #2 (Recommendations For The Authors):

The data is only appropriate a specialized journal and is limited to the Golgi analysis of neurons within the inferior olivary complex of the elephant. This reviewer considers that the remainder of the work is speculation and that the paper in its current version is not salvageable.

Comment: Rather than suggesting changes, the referee makes it clear that the referee does not want to see our paper published. We think this desire to reject is not rooted in a lack of quality of our work. In fact, we did an immense amount of work (detailed cytoarchitectonic analysis of six (now seven) elephant brainstems rather than one as in the case of our predecessors), cell counts, and X-ray tomography. Instead, we think the problem is rooted in the fact that we contradict the referee. To us, such suppression of diverging opinions – provided they are backed up with data – is a scientifically deeply unhealthy attitude. Science lives from the debate and this is why we did not exclude any referees even though we knew that our results do not align with the views of all of the few actors in the field.

Changes: We think the novel eLife publishing scheme was developed to prevent such abuse. We look forward to having our data published along with the harsh comments of the referee. The readers and subsequent scientific work will determine who’s right and who’s wrong.

In order to convince readers of the grand changes to the organization of the brainstem in a species suggested by the authors the data presented needs to be supported. It is not.

Comment: Again, this looks to us like more of the ‘total-rejection-commentary’ than like an actual recommendation.

Changes: None.

The protocol for the cytochrome oxidase histochemistry is not available in the locations indicated by the authors, and it is very necessary to provide this, as I fully believe that the staining obtained is not real, given the state of the tissue used.

Comment: We apologize again for not including the necessary details on our cytochrome-oxidase staining.

From these comments (and the initial comments above) it appears that the referee is uncertain about the validity of cytochrome-oxidase staining. We (M.B., the senior author) have been doing this particular stain for approximately three decades. The referee being unfamiliar with cytochrome-oxidase staining is fine, but we can’t comprehend how the referee then comes to the ‘full belief’ that our staining patterns are ‘not real’ when the visual evidence indicates the opposite. We feel the referee does not want to believe our data.

From hundreds of permutations, we can assure the referee that cytochrome-oxidase staining can go wrong in many ways. The most common failure outcome in elephants is a uniform light brown stain after hours or days of the cytochrome-oxidase reaction. This outcome is closely associated with long ≥2 days post-mortem/fixation times and reflects the quenching of cytochrome-oxidases by fixation. Interestingly, cytochrome-oxidase staining in elephant brains is distinctly more sensitive to quenching by fixation than cytochrome-oxidase staining in rodent brains. Another, more rare failure of cytochrome-oxidase staining comes as entirely white or barely colored sections; this outcome is usually associated with a bad reagent (most commonly old DAB, but occasionally also old or bad catalase, in case you are using a staining protocol with catalase). Another nasty cytochrome-oxidase staining outcome is smeary all-black sections. In this case, a black precipitate sticks to sections and screws up the staining (filtering and more gradual heating of the staining solution usually solve this problem). Thus, you can get uniformly white, uniformly light brown, and smeary black sections as cytochrome-oxidase staining failures. What you never get from cytochrome-oxidase staining as an artifact are sections with a strong brown to lighter brown differential contrast. All sections with strong brown to lighter brown differential contrast (staining successes) show one and the same staining pattern in a given brain area, i.e., brownish barrels in the rodent cortex, brownish barrelettes (trigeminal nuclei) in the rodent brainstem, brownish putative trunk modules/inferior olives (if we believe the referee) in the elephant brainstem. Cytochrome-oxidase reactivity is in this regard remarkably different from antibody staining. In antibody staining you can get all kinds of interesting differential contrast staining patterns, which mean nothing. Such differential contrast artifacts in antibody staining arise as a result of insufficient primary antibody specificity, the secondary antibody binding non-specifically, and of what have you not reasons. The reason that the brown differential contrast of cytochrome-oxidase reaction is pretty much fool-proof, relates to the histochemical staining mechanism, which is based on the supply of specific substrates to a universal mitochondrial enzyme. The ability to reveal mitochondrial metabolism and the universal and ‘fool-proof’ staining qualities make the cytochrome-oxidase reactivity a fantastic tool for comparative neuroscience, where you always struggle with insufficient information about antigen reactivity.

We also note that the contrast of cytochrome-oxidase reactivity seen in the elephant brainstem is spectacular. As the Referee can see in our Figure 1C we observe a dark brown color in the putative trunk module, with the rest of the brain being close to white. Such striking cytochrome-oxidase reactivity contrast has been observed only very rarely in neuroanatomy: (i) In the rest of the elephant brain (brainstem, thalamus cortex) we did not observe as striking contrast as in the putative trunk module (the inferior olive according to the referee). (ii) In decades of work with rodents, we have rarely seen such differential activity. For example, cortical whisker-barrels (a classic CO-staining target) in rodents usually come out as dark brown against a light brown background.

What all of this commentary means is that patterns revealed by differential cytochrome-oxidase staining in the elephant brain stem are real.

Changes: We added details on our cytochrome-oxidase reactivity staining protocol and commented on cytochrome-oxidase reactivity in the elephant brain in general.

The authors need to recognize that the work done in Africa on elephant brains is of high quality and should not be blithely dismissed by the authors - this stinks of past colonial "glory", especially as the primary author on these papers is an African female.

Comment: The referee notes that we unfairly dismiss the work of African scientists and that our paper reflects a continuation of our horrific colonial past because we contradict the work of an African woman. We think such commentary is meant to be insulting and prefer to return to the scientific discourse. We are staunch supporters of diversity in science. It is simply untrue, that we do not acknowledge African scientists or the excellent work done in Africa on elephant brains. For example, we cite no less than four papers from the Manger group. We refer countless times in the manuscript to these papers, because these papers are highly relevant to our work. We indeed disagree with two anatomical assignments made by Maseko et al., 2013. Such differences should not be overrated, however. As we noted before, such differences relate to only 2 out of 147 anatomical assignments made by these authors. More generally, discussing and even contradicting papers is the appropriate way to acknowledge scientists. We already expressed we greatly admire the pioneering work of the Manger group. In our view, the perfusion of elephants in the field is a landmark experiment in comparative neuroanatomy. We closely work with colleagues in Africa and find them fantastic collaborators. When the referee is accusing us of contradicting the work of an African woman, the referee is unfairly and wrongly accusing us of attacking a scientist’s identity. More generally, we feel the discussion should focus on the data presented.

Changes: None.

In addition, perfusing elephants in the field with paraformaldehyde shortly after death is not a problem "partially solved" when it comes to collecting elephant tissue (n.b., with the right tools the brain of the elephant can be removed in under 2 hours). It means the problem IS solved. This is evidenced by the quality of the basic anatomical, immuno-, and Golgi-staining of the elephant tissue collected in Africa.

Comment: This is not a recommendation. We repeat: In our view, the perfusion of elephants in the field by the Manger group is a landmark experiment in comparative neuroanatomy. Apart, from that, we think the referee got our ‘partially solved comment’ the wrong way. It is perhaps worthwhile to recall the context of this quote. We first describe the numerous limitations of our elephant material; admitting these limitations is about honesty. Then, we wanted to acknowledge previous authors who either paved the way for elephant neuroanatomy (Shoshani) or did a better job than we did (Manger; see the above landmark experiment). These citations were meant as an appreciation of our predecessors’ work and by far not meant to diminish their work. Why did we say that the problems of dealing with elephant material are only partially solved? Because elephant neuroanatomy is hard and the problems associated with it are by no means solved. Many previous studies rely on single specimen and our possibilities of accessing, removing, processing, and preserving elephant brains are limited and inferior to the conditions elsewhere. Doing a mouse brain is orders of magnitude easier than doing an elephant brain (because the problems of doing mouse anatomy are largely solved), yet it is hard to publish a paper with six elephant brains because the referees expect evidence at least half as good as what you get in mice.

Changes: We replaced the ‘partially solved’ sentence.

The authors need to give credit where credit is due - the elephant cerebellum is clearly at the core of controlling trunk movement, and as much as primary sensory and final stage motor processing is important, the complexity required for the neural programs needed to move the trunk either voluntarily or in response to stimuli, is being achieved by the cerebellum. The inferior olive is part of this circuit and is accordingly larger than one would expect.

Comment: We think it is very much possible that the elephant cerebellum is important in trunk control.

Changes: We added a reference to the elephant cerebellum in the introduction of our manuscript.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Elephants have elaborate trunk skills and large, but poorly understood brains. Here we study trunk representations in elephant trigeminal nuclei, which form large protrusions on the ventral brainstem. Dense vascularization and intense cytochrome-oxidase reactivity distinguish several elongated putative trunk modules, which repeat in the anterior-posterior direction; our analysis focuses on the most anterior and largest of the units, the putative nucleus principalis trunk module. Module neuron density is low and glia outnumbers neurons by ∼108:1. Dendritic trees are elongated along the axis of axon bundles (myelin stripes) transversing the trunk module. Furthermore, synchrotron X-ray phase contrast tomography suggests myelin-stripe-axons transverse the trunk module. We show a remarkable correspondence of trunk module myelin stripes and trunk folds. Myelin stripes show little relation to trigeminal neurons and stripe-axons appear to often go nowhere; these observations suggest to the possibility that myelin-stripes might serve to separate trunk-fold domains rather than to connect neurons. The myelin-stripes-to-folds mapping allowed to determine neural magnification factors, which changed from 1000:1 proximally to 5:1 in the trunk finger. Asian elephants have fewer (∼640,000) trunk-module neurons than Africans (∼740,000) and show enlarged representations of trunk parts involved in object wrapping. We conclude the elephant trigeminal trunk module is exquisitely organized into trunk-fold-related units.

Revised version: Elephants have elaborate trunk skills and large, but poorly understood brains. Here we study trunk representations in elephant trigeminal nuclei, which form large protrusions on the ventral brainstem. Dense vascularization and intense cytochrome-oxidase reactivity distinguish several elongated putative trunk modules, which repeat in the anterior-posterior direction; our analysis focuses on the most anterior and largest of the units, the putative nucleus principalis trunk module. Module neuron density is low and non-neural cells outnumber neurons by ∼108:1. Dendritic trees are elongated along the axis of axon bundles (myelin stripes) transversing the trunk module. Furthermore, synchrotron X-ray phase contrast tomography suggests myelin-stripe-axons transverse the trunk module. We show a remarkable correspondence of trunk module myelin stripes and trunk folds. Myelin stripes show little relation to trigeminal neurons and stripe-axons appear to often go nowhere; these observations suggest the possibility that myelin stripes might serve to separate trunk-fold domains rather than to connect neurons. The myelin-stripes-to-folds mapping allowed to determine neural magnification factors, which changed from 1000:1 proximally to 5:1 in the trunk finger. Asian elephants have fewer (∼640,000) trunk-module neurons than Africans (∼740,000) and show enlarged representations of trunk parts involved in object wrapping. We conclude the elephant trigeminal trunk module is exquisitely organized into trunk-fold-related units.

Introduction

Earlier version: Elephants are the largest extant terrestrial animals and rely on their trunks to acquire huge amounts of food. The trunk is a fusion organ of the nose and upper lip. Nose lip fusion occurs in the fourth month of elephant fetal development (Fischer & Trautmann, 1987; Schulz et al., in prep). The trunk is an immensely muscular structure (Cuvier and Dumeril, 1838; Shoshani, 1982) that contains about 90,000 muscle fascicles (Longren et al., 2023). Not surprisingly, the trunk’s prime motor control structure, the facial nucleus, is very large (Maseko et al., 2013) and shows an elaborate cellular architecture (Kaufmann et al., 2022). Trunks have prominent folds that differ between African (Loxodonta africana) and Asian (Elephas maximus) elephants (Schulz et al., in prep). Interestingly, the object-grasping behavior of African and Asian elephants differs considerably. Asian elephants have a single trunk finger and tend to wrap their trunk around objects, whereas African elephants prefer pinching objects with their two trunk fingers (Racine, 1980). Elephants are skillful with their trunks (Shoshani, 1992), have a high tactile sensitivity (Dehnhardt, Friese, and Sachser, 1997), and even acquire dexterous manipulative behaviors such as banana peeling (Kaufmann et al., 2023). Tactile feedback is of great significance for trunk behaviors, because elephants have only limited visual abilities. Recently, Deiringer et al. (2023) investigated trunk whiskers and observed use-dependent whisker lateralization, dense whisker arrays on the trunk tip and the ventral trunk, and marked whisker differences between African and Asian elephants. The sensory periphery of elephant trunks was investigated in a landmark study by Rasmussen & Munger (1996), who described dense innervation patterns in the elephant fingertip. The whole elephant trunk is massively innervated by large trigeminal ganglia (Sprinz, 1952; Purkart et al., 2022).

Revised version: Elephants are the largest extant terrestrial animals and rely on their trunks to acquire huge amounts of food. The trunk is a fusion organ of the nose and upper lip. Nose lip fusion occurs in the fourth month of elephant fetal development (Fischer & Trautmann, 1987; Schulz et al., 2023). The trunk is an immensely muscular structure (Cuvier, 1796; Cuvier and Dumeril, 1838; Shoshani, 1982) that contains about 90,000 muscle fascicles (Longren et al., 2023). Not surprisingly, the trunk’s prime motor control structure, the facial nucleus, is very large (Maseko et al., 2013) and shows an elaborate cellular architecture (Kaufmann et al., 2022). Trunks have prominent folds that differ between African (Loxodonta africana) and Asian (Elephas maximus) elephants (Schulz et al., in prep). Interestingly, the object-grasping behavior of African and Asian elephants differs considerably. Asian elephants have a single trunk finger and tend to wrap their trunk around objects, whereas African elephants prefer pinching objects with their two trunk fingers (Racine, 1980). Elephants are skillful with their trunks (Shoshani, 1992), have a high tactile sensitivity (Dehnhardt, Friese, and Sachser, 1997), and even acquire dexterous manipulative behaviors such as banana peeling (Kaufmann et al., 2023). Tactile feedback is of great significance for trunk behaviors because elephants have only limited visual abilities. Recently, Deiringer et al. (2023) investigated trunk whiskers and observed use-dependent whisker lateralization, dense whisker arrays on the trunk tip and the ventral trunk, and marked whisker differences between African and Asian elephants. The sensory periphery of elephant trunks was investigated in a landmark study by Rasmussen & Munger (1996), who described dense innervation patterns in the elephant fingertip. The whole elephant trunk is massively innervated by large trigeminal ganglia (Sprinz, 1952; Purkart et al., 2022). The cerebellum of elephants is very large both in relative and absolute terms (Maseko et al., 2012) and it is believed to be an important part of the control of the trunk movement, as a sensory and motor processing area (Maseko et al., (2013). The turning point in the investigation of the mammalian trigeminal system was the description of the cortical whisker barrels by Woolsey & Van der Loss (1970 and this work informed our approach to the elephant brainstem. The recognition of the cortical barrel pattern has led to thousands of follow-up studies, which included the discovery of whisker-related thalamic so-called barreloids (Van der Loos, 1976) and whisker-related units in the trigeminal brainstem, so-called barrelettes (Belford & Killackey, 1979; Ma, 1991). As shown by Belford & Killackey (1979) and Ma (1991) the brainstem contains several topographic trigeminal representations, which repeat in anterior to posterior direction. Specifically, these studies identified the most anterior one as the largest sensory trigeminal representation (the nucleus principalis) and several smaller more posterior trigeminal sensory nuclei. We will adopt the trigeminal terminology established by these authors (Belford & Killackey, 1979; Ma, 199).

Removed in the revised version: The turning point in the investigation of the mammalian trigeminal system has been the description of the cortical whisker barrels by Woolsey & Van der Loss (1970 and this work informed our approach to the elephant brainstem. The recognition of the cortical barrel pattern has led to thousands of follow-up studies, which included the discovery of whisker-related thalamic so-called barreloids (Van der Loos, 1976) and whisker-related units in the trigeminal brainstem, so-called barrelettes (Belford & Killackey, 1979; Ma, 1991). As shown by Belford & Killackey (1979) and Ma (1991) the brainstem contains several topographic trigeminal representations, which repeat in anterior to posterior direction. Specifically, these studies identified the most anterior one as the largest sensory trigeminal representation (the nucleus principalis) and several smaller more posterior trigeminal sensory nuclei. We will adopt the trigeminal terminology established by these authors (Belford & Killackey, 1979; Ma, 199).

A variety of excellent studies have investigated the cellular statistics of elephant brains and indicated elephant brains are not simply scaled-up mouse brains. Specifically, investigators found much lower neuronal densities in elephant brains than in rodents (Haug, 1987; Herculano-Houzel et al., 2014). A prominent difference between small and large brains is the increased amounts of white matter in larger brains. Myelin sheaths, which give white matter its whitish shine, and the enwrapped axons are usually thought to form a supply and connectivity system, an idea we will question in our study.

Earlier version: We pursued the following questions: (i) Can candidate regions for the elephant trigeminal trunk representation be identified? (ii) If yes, can multiple sensory trigeminal nuclei be identified as in other mammals? (iii) What is the neuroanatomical structure of the elephant brainstem trunk representations? (iv) Do the elaborate myelin structures in the elephant trigeminal nuclei form an axonal supply system? (v) How does the organization of elephant brainstem trunk representations relate to the differential trunk morphology and grasping behavior of African and Asian elephants? We tentatively identified an elephant brainstem trunk module characterized by intense metabolism and vascularization. The putative trunk module contains a myelin map of trunk folds. The myelin map allows precise mapping of the neural topography of the trunk representation and reveals species differences between African and Asian elephants.

Revised version: We pursued the following questions: (i) Can candidate regions for the elephant trigeminal trunk representation be identified? (ii) If yes, can multiple sensory trigeminal nuclei be identified as in other mammals? (iii) What is the neuroanatomical structure of the elephant brainstem trunk representations? (iv) Do the elaborate myelin structures in the elephant trigeminal nuclei form an axonal supply system? (v) How does the organization of elephant brainstem trunk representations relate to the differential trunk morphology and grasping behavior of African and Asian elephants? We tentatively identified an elephant brainstem trunk module characterized by intense metabolism and vascularization. The putative trunk module contains a myelin map of trunk folds. This myelin map allows precise mapping of the neural topography of the trunk representation and reveals species differences between African and Asian elephants.

Results

Overview

Determining trigeminal representations in elephants is challenging because invasive recordings or invasive viral tracing methods cannot be applied. We proceeded to build a hypothesis on the elephant trigeminal brainstem trunk in six steps. First, we identified a candidate module for the brainstem trunk representation. Second, we showed that this module architecture repeats in the anterior-posterior direction in the elephant brainstem. Third, we characterized the cellular organization of this putative trunk module. Fourth, we documented a close correspondence between the myeloarchitecture of this module and the folds of the elephant trunk. Fifth, we applied synchrotron X-ray tomography to assess the microscopic architecture of myelin stripes. Sixth, we showed that species-specific differences in trunk structure have correlates in the putative trunk module.

A metabolically highly active, strongly vascularized putative trunk module

Earlier version: The brain of the Asian elephant cow Burma is shown in Figure 1A. In rodents the sensory trigeminal nuclei are observed posterior to the pons. As shown in Figure 1B, in a ventral view of the brain stem of Burma, about 1 cm posterior to the (large) pons, a pair of large bumps is obvious on the ventral brainstem surface. By size and pronounced protrusion, these bumps on the ventral brainstem distinguish elephant brains from those of other mammals. Much smaller bumps are seen in a similar position in the human brain, where they contain the inferior olive. Accordingly, the bumps of the elephant brain have been referred to as the elephant inferior olive (Shoshani et al., 2006; Maseko et al., 2013), but our investigation did not support this idea. We sectioned the elephant medulla and stained sections for cytochrome oxidase reactivity, a mitochondrial enzyme, the activity of which is closely related to tissue energy consumption. Trigeminal nuclei tend to show intense activity in cytochrome oxidase reactivity (Belford & Killackey, 1979; Ma, 1991). A cytochrome oxidase-stained coronal section through the bump of the Asian elephant bull Raj is shown in Figure 1C. We found that the bump contained the most intense cytochrome oxidase reactivity in the elephant brainstem and (to the extent that we performed such staining in other brain regions) the rest of the elephant brain. Three cytochrome-reactive modules (a putative trunk module, a putative nostril module, and a putative lower lip/jaw) are obvious, the largest of which we refer to as a putative trunk module (Figure 1D). The putative trunk module is elongated and we hypothesize that a particularly intensely cytochrome oxidase reactive region at the ventrolateral pole of the module corresponds to the dorsal finger representation (Figure 1C, D). We provide a detailed justification for our assignments of a putative trunk module, a putative nostril module, and a putative lower lip/jaw trigeminal module in Figure 2. We also studied brain stem sections in the African elephant (Figure 1E) and identified a similar putative trunk module there. We investigated the vascularization of this module, which was evident from the cytochrome oxidase reactivity of erythrocytes in blood vessels of our non-perfused elephant brains (Figure 1F) and found that the trunk module stands out from the rest of the brainstem (Figure 1G). Specifically, it contains about twice as many blood vessels per volume as the remainder of the brainstem (Figure 1H). In parasagittal sections, the putative trunk module had a compact appearance much like the trigeminal nuclei of other mammals (Figure 1I). In parasagittal sections lateral to the putative trunk module we observed a nucleus with a very distinct banded cellular appearance (Figure 1J), a cellular architecture characteristic of the inferior olive of other mammals (Brodal et al. 1980). We conclude that the elephant brainstem contains a large, highly vascularized, and highly cytochrome-oxidase reactive elongated putative trunk module.

Revised version: The brain of the Asian elephant cow Burma is shown in Figure 1A. In rodents, the sensory trigeminal nuclei are observed posterior to the pons. As shown in Figure 1B, in a ventral view of the brain stem of Burma, about 1 cm posterior to the (large) pons, a pair of large bumps is obvious on the ventral brainstem surface. By size and pronounced protrusion, these bumps on the ventral brainstem distinguish elephant brains from those of other mammals. Much smaller bumps are seen in a similar position in the human brain, where they contain the inferior olive. Accordingly, the bumps of the elephant brain have been referred to as the elephant inferior olive (Shoshani et al., 2006; Maseko et al., 2013; Verhaart and Kramer, 1958; Verhaart 1962), but our investigation did not support this idea. We sectioned the elephant medulla and stained sections for cytochrome oxidase reactivity, a mitochondrial enzyme, the activity of which is closely related to tissue energy consumption. Trigeminal nuclei tend to show intense activity in cytochrome oxidase reactivity (Belford & Killackey, 1979; Ma, 1991). A cytochrome oxidase-stained coronal section through the bump of the Asian elephant bull Raj is shown in Figure 1C. We found that the bump contained the most intense cytochrome oxidase reactivity in the elephant brainstem and (to the extent that we performed such staining in other brain regions) the rest of the elephant brain. Three cytochrome-reactive modules (a putative trunk module, a putative nostril module, and a putative lower lip/jaw) are obvious, the largest of which we refer to as a putative trunk module (Figure 1D). The putative trunk module is elongated and we hypothesize that a particularly intensely cytochrome oxidase reactive region at the ventrolateral pole of the module corresponds to the dorsal finger representation (Figure 1C, D). We provide a detailed justification for our assignments of a putative trunk module, a putative nostril module, and a putative lower lip/jaw trigeminal module in Figure 2. We also studied brain stem sections in the African elephant (Figure 1E) and identified a similar putative trunk module there. We investigated the vascularization of this module, which was evident from the cytochrome oxidase reactivity of erythrocytes in blood vessels of our non-perfused elephant brains (Figure 1F), and found that the trunk module stands out from the rest of the brainstem (Figure 1G). Specifically, it contains about twice as many blood vessels per volume as the remainder of the brainstem (Figure 1H). In parasagittal sections, the putative trunk module had a compact appearance much like the trigeminal nuclei of other mammals (Figure 1I). In parasagittal sections lateral to the putative trunk module we observed a nucleus with a very distinct banded cellular appearance (Figure 1J), a cellular architecture characteristic of the inferior olive of other mammals (Brodal et al. 1980). We conclude that the elephant brainstem contains a large, highly vascularized, and highly cytochrome-oxidase reactive elongated putative trunk module.

Appearance, metabolism, blood supply, and identification of the elephant trigeminal nuclei

A, ventral view of the brain of Burma, a 52-year-old female Asian elephant.

B, ventral view of the brainstem of Burma. Note the large pons (upper part of the photograph). Posterior to the pons a pair of prominent bumps are visible, the putative trigeminal nuclei (TN) of the elephant. The trigeminal nuclei bumps are observed, wherein in humans and other mammals the inferior olive is found and was referred to as inferior olive by previous authors (Shoshani et al., 2006; Maseko et al., 2013; Rasenberger, 2019). The inferior olive can be identified at an unusual lateral position in elephants, however (IO).

C, 60 µm coronal section through the trigeminal nuclei of Raj, a four-year-old elephant bull, stained for cytochrome-oxidase reactivity, a mitochondrial enzyme, the reactivity of which reflects constitutive metabolic activity. The trigeminal nuclei show some of the strongest cytochrome-oxidase reactivity (indicated by the brown color) in the elephant brain and strong cytochrome-oxidase reactivity is typical for the trigeminal nuclei of many mammals.

D, a drawing of putative trigeminal subnuclei stained in C. Note the compact shape of the trunk module, which is unlike the inferior olive of mammals.

E, drawing of a coronal section stained for cytochrome-oxidase reactivity through the brainstem of Indra a female African elephant, borders of nuclei are outlined.

F, upper, micrograph of cytochrome-oxidase reactivity in the putative trunk fingertip representation. Lower, drawing of cytochrome-oxidase reactive erythrocytes in blood vessels. Note the high density of vessels inside of the fingertip but not in the surrounding tissue.

G, drawing of the entire brainstem section. The putative trunk module stands out from the rest of the brainstem in terms of vessel density.

H, quantification of blood vessel length in various parts of the brainstem. Note that not all vessels contain erythrocytes and are stained, i.e., measures of blood vessel length are lower bound estimates.

I, Brightfield micrograph of a parasaggital section through the trigeminal trunk module. Note the compact cellular architecture.

J, Brightfield micrograph of a parasaggital section through the inferior olive. Note the banded cellular architecture that is characteristic of the mammalian inferior olive.

L = lateral; V = ventral; P = posterior.

Added in the revised version: The putative inferior olive but not the putative trigeminal nucleus contains peripherin-positive axon bundles (presumptive climbing fibers).

Added in the revised version: A, Overview picture of a brainstem section stained with anti-peripherin-antibodies (white color). Anti-peripherin-antibodies stain climbing fibers in a wide variety of mammals. The section comes from the posterior brainstem of African elephant cow Bibi; in this posterior region, both putative inferior olive and trigeminal nucleus are visible. Note the bright staining of the dorsolateral nucleus, the putative inferior olive according to Reveyaz et al., and the trigeminal nucleus according to Maseko et al., 2013.

Added in the revised version: B, High magnification view of the dorsolateral nucleus (corresponding to the upper red rectangle in A). Anti-peripherin-positive axon bundles (putative climbing fibers) are seen in support of the inferior olive hypothesis of Reveyaz et al.

Added in the revised version: C, High magnification view of the ventromedial nucleus (corresponding to the lower red rectangle in A). The ventromedial nucleus is weakly positive for peripherin but contains no anti-peripherin-positive axon bundles (i.e. no putative climbing fibers) in support of the trigeminal nucleus hypothesis of Reveyaz et al. Note that myelin stripes – weakly visible as dark omissions – are clearly anti-peripherin-negative.

Added in the revised version: Previous work on the elephant brainstem (Shoshani et al., 2006; Maseko et al., 2013; Verhaart and Kramer, 1958; Verhaart 1962) suggested that the structure we assigned as trigeminal nucleus is in fact the inferior olive; this matter is also discussed in depth in our earlier discussions with the referees, which are published along with our article. We performed additional peripherin-antibody staining to assess the possibility that the structure we assigned as trigeminal nucleus corresponds to the inferior olive. To differentiate between the trigeminal nuclei and the inferior olive, we used a climbing fiber antibody staining. Peripherin is a cytoskeletal protein that is found in peripheral nerves and climbing fibers. Specifically, climbing fibers of various species (mouse, rabbit, pig, cow, and human; Errante et al 1998) stain intensely with peripherin- antibodies. As can be seen in Figure 2, we observe peripherin-reactivity in axonal bundles (i.e. in putative climbing fibers), in what we think is the inferior olive (Figure 2B). We also observe some peripherin-reactivity in what we think is the trigeminal nucleus, but not the distinct and strong labeling of axonal bundles (Figure 2C). This lack of peripherin-reactive axon bundles suggests, that there are no climbing fibers, in what was previously thought as the inferior olive of the elephants. Peripherin-staining is consistent with the idea that the ventral nuclei described by us correspond to the trigeminal nuclei of the elephant and the more lateral nuclei, previously described as trigeminal nuclei by Maseko et al., 2013 correspond to the inferior olive.

Added in the revised version: Trigeminal nuclei in coronal and horizontal sections of African elephant brainstem

Added in the revised version: We found that putatively trunk-related trigeminal modules repeat at least two and probably four times in the anterior-posterior direction in the elephant brainstem. All these repeating modules had a higher cytochrome oxidase reactivity and a higher cell density than surrounding brainstem structures. Such repeats of trigeminal representations in the anterior-posterior direction are also seen in other mammals (Belford & Killackey, 1979; Ma, 1991). We refer to these modules with the same terminology as established in rodents. As in other mammals, we found the most anterior representation to be larger than the others, and we refer to this representation as nucleus principalis (Pr5, which stands for principal trigeminal nucleus). Figure 3A shows a Nissl-stained coronal section through the principalis trigeminal modules. In Figure 3B, C we provide a color-coded putative assignment of principalis modules. We assigned the large (grey) module to the trunk, because of its cytochrome oxidase reactivity, its elongation, and its extraordinary size. We assigned the elongated (red/pink) module to the nostril for the following reasons: 1. its unusually (among brainstem modules) thin tube-like appearance. 2. The widening towards the tip of the putative trunk module. 3. The cellular continuity with the mouth opening of the putative trunk module. 4. The topographic relationship with the lower jaw module, which matched the topography of the elephant head (Figure 3C). 5. The fact that this module had the same length as the trunk module. 6. The fact that there were no indications of a nostril module inside the putative trunk module, (where we initially expected a nostril representation). Our reasons for assigning the compact (blue) module to the lower jaw were its shape (Figure 3A-C) and topographic position. Next, we provide an overview of the arrangement of trigeminal modules in horizontal sections (Figure 3D-F, proceeding from dorsal to ventral). At the dorsal level (Figure 3D) only two trunk modules (TM), can be recognized. These are Pr5 and Sp5o, which stands for spinal trigeminal nucleus pars oralis TM, directly posterior to the Pr5. The cell density is low, the modules barely stand out from the surroundings and we think that at the dorsal level proximal trunk parts are represented. At the midlevel (Figure 3E) four repeating trunk modules (Pr5 TM, Sp5o TM, and Sp5i, which stands for spinal trigeminal nucleus pars interpolaris, and Sp5c, which stands for spinal trigeminal nucleus pars caudalis TM) can be recognized. We did not investigate facial representations other than the trunk module. The identification of the Sp5i and Sp5c trunk modules is only tentative at this point. The analysis of horizontal and parasagittal sections pointed to a mirror image-like arrangement of these modules. The cell density is higher at midlevel (Figure 3E). At the ventral level (Figure 3F) only two trunk modules (Pr5 and Sp5o TM, directly posterior to the principalis) can be recognized. In this section, the mirror image- like arrangement of the Pr5 TM and the Sp5o TM is evident. The cell density is very high and we think that the trunk tip is represented here. We conclude that repeating trigeminal trunk modules can be recognized in the elephant brainstem.

Overview of trigeminal nuclei in coronal and horizontal sections of African elephant cow Indra

A, micrograph of Nissl-stained coronal 60-µm-section through the right-hemispheric principalis trunk module of elephant cow Indra. The principalis is the most anterior and by far the largest trigeminal representation in elephants.

B, drawing of color-coded trigeminal modules belonging to the principalis nucleus.

C, drawing of an African elephant head with different facial structures color-coded according to the trigeminal modules they putatively correspond to in A, B.

D, left micrograph of Nissl-stained horizontal section through the left-hemispheric trigeminal nuclei of elephant cow Indra. The section is positioned at a dorsal level of the trigeminal nuclei. Right, drawing of trigeminal modules shown in the left micrograph; for the principalis (Pr5) module the same color code as in B, and C has been used. Sp5o refers to the trunk module directly posterior to the principalis nucleus. The facial nucleus and the putative inferior olive were also identified.

E, conventions as in D. Horizontal section through the mid-level of the trigeminal nuclei. Sp5i TM and Sp5c TM, refer to putative trunk modules posterior to the Sp5o TM.

F, conventions as in D. Horizontal section through the ventral level of the trigeminal nuclei.

Pr5, principal trigeminal nucleus; Sp5o, spinal trigeminal nucleus pars oralis; Sp5i, spinal trigeminal nucleus pars interpolaris; Sp5c, spinal trigeminal nucleus pars caudalis; TM = trunk module; A = anterior; L = lateral; V = ventral.

Removed in the revised version: Trigeminal nuclei in coronal and horizontal sections of African elephant brainstem

Removed in the revised version: We found that putatively trunk-related trigeminal modules repeat at least two and probably four times in the anterior-posterior direction in the elephant brainstem. All these repeating modules had a higher cytochrome oxidase reactivity and a higher cell density than surrounding brainstem structures. Such repeats of trigeminal representations in the anterior-posterior direction are also seen in other mammals (Belford & Killackey, 1979; Ma, 1991). We refer to these modules with the same terminology as established in rodents. As in other mammals, we found the most anterior representation to be larger than the others, and we refer to this representation as nucleus principalis (Pr5, which stands for principal trigeminal nucleus). Figure 2A shows a Nissl-stained coronal section through the principalis trigeminal modules. In Figure 2B, C we provide a color-coded putative assignment of principalis modules. We assigned the large (grey) module to the trunk, because of its cytochrome oxidase reactivity, its elongation, and its extraordinary size. We assigned the elongated (red/pink) module to the nostril for the following reasons: 1. its unusually (among brainstem modules) thin tube-like appearance. 2. The widening towards the tip of the putative trunk module. 3. The cellular continuity with the mouth opening of the putative trunk module. 4. The topographic relationship with the lower jaw module, which matched the topography of the elephant head (Figure 2C). 5. The fact that this module had the same length as the trunk module. 6. The fact that there were no indications of a nostril module inside the putative trunk module, (where we initially expected a nostril representation). Our reasons for assigning the compact (blue) module to the lower jaw were its shape (Figure 2A-C) and topographic position. Next, we provide an overview of the arrangement of trigeminal modules in horizontal sections (Figure 2D-F, proceeding from dorsal to ventral). At the dorsal level (Figure 2D) only two trunk modules (TM), can be recognized. These are Pr5 and Sp5o, which stands for spinal trigeminal nucleus pars oralis TM, directly posterior to the Pr5. The cell density is low, the modules barely stand out from the surroundings and we think that at the dorsal level proximal trunk parts are represented. At the midlevel (Figure 2E) four repeating trunk modules (Pr5 TM, Sp5o TM, and Sp5i, which stands for spinal trigeminal nucleus pars interpolaris, and Sp5c, which stands for spinal trigeminal nucleus pars caudalis TM) can be recognized. We did not investigate facial representations other the trunk module. The identification of the Sp5i and Sp5c trunk modules is only tentative at this point. The analysis of horizontal and parasagittal sections pointed to a mirror image-like arrangement of these modules. The cell density is higher at midlevel (Figure 2E). At the ventral level (Figure 2F) only two trunk modules (Pr5 and Sp5o TM, directly posterior to the principalis) can be recognized. In this section, the mirror image-like arrangement of the Pr5 TM and the Sp5o TM is evident. The cell density is very high and we think that the trunk tip is represented here. We conclude that repeating trigeminal trunk modules can be recognized in the elephant brainstem.

The cellular architecture of the putative principalis trunk module

Earlier version: We studied the module’s cellular organization by Golgi stains of the putative principalis trunk module of the Asian elephant Raj (Figure 3A). Golgi stains identified two prominent neuronal elements. First, we observed bundles of large diameter (3-15 µm) axons, which run orthogonal to the module’s main axis in coronal sections (Figure 3B arrows). These axon bundles correspond to the myelin stripes described in Figure 4. Somato-dendritically stained neurons were the second neuronal element identified by Golgi stains (Figure 3C). The density of Golgi-stained neurons was very low. We reconstructed neurons using a Neurolucida system and superimposed 47 neurons (from three adjacent sections) in Figure 3D. The most abundant cells in the trigeminal nucleus were putative astrocytes (Figure 3E, green). We distinguished two types of neurons, putative principal neurons with large somata and branched dendrites (n = 41; black in Figure 3E) and other neurons, putative interneurons, with small somata and unbranched dendrites (n = 6; red in Figure 3E). Putative astrocytes, principal neurons, and interneurons differed markedly in their morphologies (Figure 3E, Table 1). The dendritic trees of neurons were elongated (Figure 3E), an observation confirmed when we prepared raw polar plots of dendritic orientation (Figure 2F, upper). We had the impression that dendritic elongation and axon bundles followed the same axis. We tested the idea that dendritic trees were aligned to myelin stripes was by rotating dendritic trees and aligning all trees according to the local axon bundle orientations. When we aligned dendritic trees this way, we observed an even stronger population polarization of dendritic trees, i.e., dendritic trees were average twofold longer along the axon bundle axis (Figure 3F, lower).

Revised version: We studied the module’s cellular organization by Golgi stains of the putative principalis trunk module of the Asian elephant Raj (Figure 4A). Golgi stains identified two prominent neuronal elements. First, we observed bundles of large diameter (3-15 µm) axons, which run orthogonal to the module’s main axis in coronal sections (Figure 4B arrows). These axon bundles correspond to the myelin stripes described in Figure 5. Somato-dendritically stained neurons were the second neuronal element identified by Golgi stains (Figure 4C). The density of Golgi-stained neurons was very low. We reconstructed neurons using a Neurolucida system and superimposed 47 neurons (from three adjacent sections) in Figure 4D. The most abundant cells in the trigeminal nucleus were putative astrocytes (Figure 4E, green). We distinguished two types of neurons, putative principal neurons with large somata and branched dendrites (n = 41; black in Figure 4E) and other neurons, putative interneurons, with small somata and unbranched dendrites (n = 6; red in Figure 4E). Putative astrocytes, principal neurons, and interneurons differed markedly in their morphologies (Figure 4E, Table 1). The dendritic trees of neurons were elongated (Figure 4E), an observation confirmed when we prepared raw polar plots of dendritic orientation (Figure 4F, upper). We had the impression that dendritic elongation and axon bundles followed the same axis. We tested the idea that dendritic trees were aligned to myelin stripes by rotating dendritic trees and aligning all trees according to the local axon bundle orientations. When we aligned dendritic trees this way, we observed an even stronger population polarization of dendritic trees, i.e., dendritic trees were average twofold longer along the axon bundle axis (Figure 4F, lower).

Cellular organization of the putative trunk module in the elephant trigeminal nuclei.

A, Golgi stained 200 µm coronal section through the trunk module of Raj, a four-year- old Asian elephant bull.

B, bundles of very thick axons are revealed by the Golgi staining (red arrows). Bundles were oriented orthogonal to the main axis of the module except for the putative finger representation, where they curved around.

C, Golgi-stained neurons are also observed albeit at a low frequency.

D, well-stained and well-preserved neurons reconstructed with a Neurolucida system. We show 47 neurons reconstructed from three adjacent coronal Golgi sections superimposed. Note the small size of the neurons relative to the module.

E, left (green), ten putative astrocytes. These small cells are the most abundant cellular element in the trigeminal nucleus. Right, four neuronal reconstructions are shown at higher magnification. Putative principal cells (three shown in black, total n = 41) had large somata and branched dendrites. A few cells (one cell shown in red, total n = 6) had small somata and unbranched dendrites. Dendritic trees were weakly polarized.

F, upper, raw polar plot of the orientation of neuronal dendritic segments (from all putative principal cells and putative interneurons, n = 47) relative to the soma confirms the common elongation of dendrites. Lower, when cells were aligned to the local axon bundle orientation an even stronger polarization of dendrites is evident.

G, antibody staining of neurons (green fluorescence, NeuN-antibody) and nuclei of all cells (blue fluorescence, DAPI) of a coronal section through the putative trunk module of Indra, a 34-year-old female African elephant. Neuron density is low.

Earlier version: H, high magnification view of the section shown in G, glia outnumbers neurons by about a hundredfold (data refer to neuron and glia counts from three elephant trigeminal nuclei).

Revised version: H, high magnification view of the section shown in G, non-neural cells outnumber neurons by about a hundredfold (data refer to neuron and non-neural cell counts from three elephant trigeminal nuclei).

I, upper, somata drawing from a Nissl stained 60 µm coronal section through the putative trunk module of Indra. Lower, cells from the medial, the putatively proximal trunk representation of the module, and the lateral, the putatively distal trunk representation. Note the soma size difference.

J, plot of soma area along the length of the module. Neurons were sequentially measured along the axis of the module. Each dot refers to one of 1159 neurons in the section; red running average (across 40 neurons) of soma area. Cells are significantly larger in lateral (putatively distal trunk representation), unpaired T-test.

FT, putative dorsal Finger Tip representation; V = ventral; L = lateral.

Trunk module myelin stripes form a precise map of trunk folds

A, a brightfield image of a freshly cut 60 µm coronal section through the center of the putative trunk module of adult female elephant Indra. Neurons are evident as small white dots and whitish myelin stripes are readily apparent even in this unstained tissue. B, a fluomyelin stain (green fluorescence) confirms that stripes contain myelin. High- resolution brightfield microscopy (not shown) and Golgi stains (Figure 2B) confirm that the stripes consist of axon bundles.

C, a line drawing (red) of myelin stripes superimposed to the micrograph shown in A.

Earlier version: D, upper, enlarged view of line drawing (red) of myelin stripes shown in C; we quantified dorsally ending, ventrally ending, and full transversal stripes. Such numbersmatch the number of trunk folds quantified in E. Lower, based on the idea of a match of trunk folds with myelin stripes one can compute magnification factors across the trunk module. Neural data (green) refer to distances between dorsally ending myelin stripes (i.e., neural distances were measured along the dorsal border (green line) of the module). Trunk distances were measured between dorsal trunk folds.

Revised version: D, upper, enlarged view of line drawing (red) of myelin stripes shown in C; we quantified dorsally ending, ventrally ending, and full transversal stripes. Such numbers match the number of trunk folds quantified in E.

Added in the revised version: Lower, based on the idea of a match of trunk folds with myelin stripes one can compute magnification factors across the trunk module. Neural data (green) refer to distances between dorsally ending myelin stripes (i.e., neural distances were measured along the dorsal border (green line) of the module). Trunk distances were measured between dorsal trunk folds.

E, drawing of the dorsal folds of the trunk of Indra.

F, the composite photograph of the dorsal trunk of Indra. We counted dorsal folds, ventral folds (not visible here), and folds that fully transversed the right trunk side (not visible here). Trunk folds match in number, orientation (typically transversal), and patterning with myelin stripes seen on the trunk module.

V = ventral; L = lateral.

Earlier version: We stained coronal sections with NeuN-antibody to identify neurons and used the DNA-stain DAPI to identify all cell nuclei (Figure 3G). Neuronal density was fairly low, but the density of non-neuronal cells was substantial (Figure 3H). In counts of individual fluorescence sections, we observed a ratio of neurons to putative glia of 1 to 80 (in one elephant). We then made more systematic counts of neurons vs. visually identified glia across three entire trunk modules from two elephants. We then observed a ratio of 1 neuron to 108 ± 24 glia. We also observed that neuron size was not homogeneous across the putative trunk module. We observed a cell size increase from proximal to distal in all putative modules (n = 12), for which we had cellular stains. To quantify cell size differences, we drew somata from a Nissl-stained section through the center of the trunk module of the African elephant cow Indra (Figure 3I). We found cell size to increase significantly from the putative proximal to the putative distal finger representation (Figure 3J). We conclude that the putative trunk module contains transversally running axon bundles and neurons, which are vastly outnumbered by glia.

Revised version: We stained coronal sections with NeuN-antibody to identify neurons and used the DNA-stain DAPI to identify all cell nuclei (Figure 4G). Neuronal density was fairly low, but the density of non-neuronal cells was substantial (Figure 4H). In counts of individual fluorescence sections, we observed a ratio of neurons to non-neural cells of 1 to 80 (in one elephant). We then made more systematic counts of neurons vs. visually identified non-neural cells across three entire trunk modules from two elephants. We then observed a ratio of 1 neuron to 108 ± 24 non-neural cells. We also observed that neuron size was not homogeneous across the putative trunk module. We observed a cell size increase from proximal to distal in all putative modules (n = 12), for which we had cellular stains. To quantify cell size differences, we drew somata from a Nissl-stained section through the center of the trunk module of the African elephant cow Indra (Figure 4I). We found cell size to increase significantly from the putative proximal to the putative distal finger representation (Figure 4J). We conclude that the putative trunk module contains transversally running axon bundles and neurons, which are vastly outnumbered by non-neural cells.

Module myelin stripes match with number, orientation, and patterning of trunk folds

Earlier version: In Nissl or cytochrome-oxidase stains, we observed prominent myelin stripes apparent as white omissions. Remarkably, entirely unstained freshly-cut coronal brainstem sections showed the clearest stripe pattern in brightfield microscopy (Figure 4A). Fluorescent stains for myelin (fluomyelin) confirmed the presence of myelin (Figure 4B). As already suggested by Golgi staining, the myelin stripes appeared to consist of large-diameter axons. The visibility of myelin stripes varied with the sectioning plane and anterior-posterior position. Myelin stripes were most obvious at the anterior-posterior center of the putative trunk module, as seen for the coronal section in Figure 4A. We investigated the myelin-stripe trunk correspondence. To this end, we made drawings of myelin stripes (Figure 4C) and compared the pattern of myelin stripes (Figure 4D upper) to the pattern of trunk folds of Indra’s trunk (Figure 4E, F). Myelin stripe and trunk fold patterns were very similar. In all trunk modules sectioned, we observed an overall match of stripe and trunk fold orientation (to the module and trunk main axis, respectively). We also observed in all modules a lack of fully transversing stripes in the putative finger region of the module, which is consistent with the lack of folds across the trunk ‘mouth’. In favorable cases, where we had brightfield images of the trunk module and had access to the elephant’s trunk, the data hinted at a 1 to 1 matching of stripes and folds. Specifically, we observed 65 myelin stripes ending on the dorsal side of the module, 46 ventrally ending myelin stripes, and 31 full transversal myelin stripes. In terms of folds, we observed 64 dorsal trunk folds (Figure 4E), 49 folds on the ventral side of the trunk, and 32 folds that fully transversed the right side of Indra’s trunk. This numeric correspondence is very suggestive and inspired a detailed mapping of trunk sensory topography (Figure 4D lower). Based on stripe-wrinkle matching, we suggest that sensory magnification increases from 1000:1 (trunk: trigeminal nucleus) in the proximal representation of the trunk module to 5:1 in the trunk-finger representation. Sectioning angle was a major factor determining match between myelin stripes and folds., i.e. we observed myelin stripes in trunk fold-like patterns in all coronally sectioned specimens. In horizontally sectioned elephant brainstem, myelin stripes were seen but could not be related to folds. In parasagitally sectioned brainstems few myelin stripes were obvious. In coronal sections, myelin stripes were most obvious in the center of the module and matched best to trunk folds. The staining method was another determinant of the match. Myelin stripes were best visible in unstained freshly cut sections with brightfield microscopy. As expected, myelin stripes were also stained positively for fluorescent myelin dyes, such as fluomyelin-green (Figure 4B) or fluomyelin-red (data not shown). Nissl or cytochrome oxidase stains were less sensitive than visualizing myelin stripes in brightfield images, i.e. not all stripes are visible in each section. While in some sections like the one shown in Figure 4D, pretty much all stripes could be mapped to trunk folds, most sections contained a few stripes that had deviating trajectories from the other stripes and these stripes could not be mapped to trunk folds. A good match of myelin stripes to folds depended also on the assessment of trunk folds. Specifically, a good match was only obtained, if we restricted folds counts to major trunk wrinkles/folds, minor trunk wrinkles appear not to be robustly represented by myelin stripes. We conclude myelin stripe patterns behave not unlike rodent cortical barrel patterns, the visibility of which also greatly depends on the staining method and sectioning angle.

Revised version: In Nissl or cytochrome-oxidase stains, we observed prominent myelin stripes apparent as white omissions. Remarkably, entirely unstained freshly-cut coronal brainstem sections showed the clearest stripe pattern in brightfield microscopy (Figure 5A). Fluorescent stains for myelin (fluomyelin) confirmed the presence of myelin (Figure 5B). As already suggested by Golgi staining, the myelin stripes appeared to consist of large-diameter axons. The visibility of myelin stripes varied with the sectioning plane and anterior-posterior position. Myelin stripes were most obvious at the anterior- posterior center of the putative trunk module, as seen for the coronal section in Figure 5A. We investigated the myelin-stripe trunk correspondence. To this end, we made drawings of myelin stripes (Figure 5C) and compared the pattern of myelin stripes (Figure 5D upper) to the pattern of trunk folds of Indra’s trunk (Figure 5E, F). Myelin stripe and trunk fold patterns were very similar. In all trunk modules sectioned, we observed an overall match of stripe and trunk fold orientation (to the module and trunk main axis, respectively). We also observed in all modules a lack of fully transversing stripes in the putative finger region of the module, which is consistent with the lack of folds across the trunk ‘mouth’. In favorable cases, where we had brightfield images of the trunk module and had access to the elephant’s trunk, the data hinted at a 1 to 1 matching of stripes and folds. Specifically, we observed 65 myelin stripes ending on the dorsal side of the module, 46 ventrally ending myelin stripes, and 31 full transversal myelin stripes. In terms of folds, we observed 64 dorsal trunk folds (Figure 5E), 49 folds on the ventral side of the trunk, and 32 folds that fully transversed the right side of Indra’s trunk. This numeric correspondence is very suggestive and inspired a detailed mapping of trunk sensory topography (Figure 5D lower). Based on stripe- wrinkle matching, we suggest that sensory magnification increases from 1000:1 (trunk: trigeminal nucleus) in the proximal representation of the trunk module to 5:1 in the trunk-finger representation. Sectioning angle was a major factor in determining the match between myelin stripes and folds, i.e. we observed myelin stripes in trunk fold- like patterns in all coronally sectioned specimens. In horizontally sectioned elephant brainstem, myelin stripes were seen but could not be related to folds. In parasagitally sectioned brainstems few myelin stripes were obvious. In coronal sections, myelin stripes were most obvious in the center of the module and matched best to trunk folds. The staining method was another determinant of the match. Myelin stripes were best visible in unstained freshly cut sections with brightfield microscopy. As expected, myelin stripes were also stained positively for fluorescent myelin dyes, such as fluomyelin-green (Figure 5B) or fluomyelin-red (data not shown). Nissl or cytochrome oxidase stains were less sensitive than visualizing myelin stripes in brightfield images, i.e. not all stripes are visible in each section. While in some sections like the one shown in Figure 5D, pretty much all stripes could be mapped to trunk folds, most sections contained a few stripes that had deviating trajectories from the other stripes and these stripes could not be mapped to trunk folds. A good match of myelin stripes to folds depended also on the assessment of trunk folds. Specifically, a good match was only obtained, if we restricted fold counts to major trunk wrinkles/folds, minor trunk wrinkles appear not to be robustly represented by myelin stripes. We conclude myelin stripe patterns behave not unlike rodent cortical barrel patterns, the visibility of which also greatly depends on the staining method and sectioning angle.

Myelin stripe architecture and the lack of a relation of stripes to trigeminal neurons

Earlier version: Their large size makes determining the architecture of myelin stripes is a challenging. To confront this challenge, we applied synchrotron-powered X-ray phase-contrast tomography of an 8 mm unstained and paraffin-embedded trigeminal nucleus tissue punch (Figure 5A,B); based on phase-contrast this methodology allows us to sample large image volumes (Figure 5C) with submicrometer (0.65 µm isotropic voxel size) resolution. Such imaging allowed us to identify myelin stripes in unstained trigeminal tissue (Figure 5D) and even enabled the reconstruction of individual large-diameter axons for several millimeters through the entire volume image (Figure 5E). As observed before with light microscopy, myelin stripes ran in the coronal plane and were about seven myelinated axons wide (maximum extent in the coronal plane; Figure 5F). Myelin stripes were circular axon bundles (Figure 5G) and were also about seven myelinated axons high (maximum extent in the anterior-posterior plane; Figure 5H). With that, we estimated that stripes are made up of 20-50 myelinated axons, unmyelinated axons could not be resolved in our analysis. The large-diameter axons, which could be followed through the X-ray tomography volume image followed an ‘all the way’ pattern (i.e. fully transversing the module). We also found that myelin stripes have a fairly consistent thickness from their dorsal to their ventral end (Figure 5I). This observation argues against the idea that myelin stripes are conventional axonal supply structures, from which axons divert off into the tissue. We also analyzed myelin stripes throughout the trunk module (Figure 5J) to understand how their thickness relates to trigeminal neuron numbers (i.e. the number of neurons between myelin stripes; Figure 5J). We observed little obvious relation between myelin stripe thickness and trigeminal neuron number. We conclude that myelin stripes have a stereotyped architecture, but show little relation to trigeminal neurons.

Revised version: Their large size makes determining the architecture of myelin stripes challenging. To confront this challenge, we applied synchrotron-powered X-ray phase-contrast tomography of an 8 mm unstained and paraffin-embedded trigeminal nucleus tissue punch (Figure 6A,B); based on phase-contrast this methodology allows us to sample large image volumes (Figure 6C) with submicrometer (0.65 µm isotropic voxel size) resolution. Such imaging allowed us to identify myelin stripes in unstained trigeminal tissue (Figure 6D) and even enabled the reconstruction of individual large-diameter axons for several millimeters through the entire volume image (Figure 6E). As observed before with light microscopy, myelin stripes ran in the coronal plane and were about seven myelinated axons wide (maximum extent in the coronal plane; Figure 6F). Myelin stripes were circular axon bundles (Figure 6G) and were also about seven myelinated axons high (maximum extent in the anterior-posterior plane; Figure 6H). With that, we estimated that stripes are made up of 20-50 myelinated axons, unmyelinated axons could not be resolved in our analysis. The large-diameter axons, which could be followed through the X-ray tomography volume image followed an ‘all the way’ pattern (i.e. fully transversing the module). We also found that myelin stripes have a fairly consistent thickness from their dorsal to their ventral end (Figure 6I). This observation argues against the idea that myelin stripes are conventional axonal supply structures, from which axons divert off into the tissue. We also analyzed myelin stripes throughout the trunk module (Figure 6J) to understand how their thickness relates to trigeminal neuron numbers (i.e. the number of neurons between myelin stripes; Figure 6J). We observed little obvious relation between myelin stripe thickness and trigeminal neuron number. We conclude that myelin stripes have a stereotyped architecture, but show little relation to trigeminal neurons.

Microscopic organization of myelin stripes and absence of a strong relation of stripes to trigeminal neurons

A, upper, schematic of the trunk module with myelin stripes (grey) of African elephant Bambi and targeting of the 8 mm tissue punch (dashed red circle). Lower, synchrotron radiation (red flash, DESY, Hamburg) was directed to the area of the punch (black box) and imaged. V = ventral; L = lateral.

B, sketch of the parallel beam setup of the GINIX endstation (P10 beamline, DESY, Hamburg). In this geometry, a dataset of the trunk module was acquired at an effective voxel size of 650nm3

C, dimensions of the imaged volume shown as a volume rendering (0.65 µm isotropic voxel size)

D, transparent image volume. Two myelin stripes were followed through the volume image (highlighted red axon bundle). Four large diameter (∼ 15 µm) axons were also reconstructed and could be followed through the entire volume image (blue).

E, axon bundles (red) and reconstructed axons (blue) in isolation.

F, image section in the coronal plane at the center of the myelin stripe. The myelin stripe (pink overlay) is readily visible, the reconstructed axon is highlighted in blue, and the bundle has a width of about 7 myelinated axons.

G, cross-section through the axon bundle (pink overlay). The virtual section is cut orthogonal to the coronal plane. The reconstructed axon is highlighted in blue.

H, image section orthogonal to the coronal plane, the virtual section is cut in an anterior-posterior direction parallel to the axon bundle (pink overlay). The myelin stripe is readily visible, the reconstructed axon is highlighted in blue, and the bundle has a height of about 7 myelinated axons.

I, left, myelin stripes. Right, measurement of the thickness (width orthogonal to the main stripes axis) along the dorsoventral axis of stripes; data are the average of ten measurements along ten myelin stripes that fully transversed the module. Stripe thickness is fairly constant, there is no evidence of stripe tapering as would be expected if axons bud off into the tissue.

J, upper, measurement of the thickness (width orthogonal to the main stripes axis) of myelin stripes across the putative trunk module; only full transversal stripes were measured. Stripe thickness is fairly constant. Lower, neuron number between full transversal myelin stripes. Neuron number varies more than 100-fold and is very low (zero) between medial stripes (in the putatively proximal trunk representation). The fact that stripe thickness changes little across the module, while neuron number between stripes changes massively argues against a relationship between stripe thickness and trigeminal neuron number.

The putative trunk module mirrors species differences in trunk folds and trunk use

Earlier version: We found that the trigeminal bumps on the ventral brainstem differ significantly between African and Asian elephants (Figure 6A). We counted neurons in the principalis trunk module and found that African elephants (740210 ± 51902, mean ± SD) had more neurons than Asian elephants (636447 ± 69729, mean ± SD) and also had a larger volume principalis trunk module (Figure 6B). Supplementary Tables 2 and 3 provide further information on our counts of the trigeminal nuclei. As noted there, the dorsal finger accounted for a large fraction (∼20%) of the trunk modules. We wondered, why brainstem bumps differed between African and Asian elephants, and therefore closely investigated the shape of trunk modules in these species. A cytochrome-oxidase stained coronal section through the trunk module of the African elephant Indra is shown in Figure 6C. Drawings of coronal sections from this trunk module and the trunk module of other African elephants are shown in Figure 6D; myelin stripes (violet) were visible as whitish omissions of the cytochrome oxidase or of Nissl stains. We also determined the length, the width, and the longitudinal position of the greatest width of the module (black line; Figure 6D). A cytochrome-oxidase-stained coronal section through the trunk module of Asian elephant Raj is shown in Figure 6E and drawings from this and other trunk modules of Asian elephants are shown in Figure 6F. Basic aspects of the module trunk were similar in African and Asian elephants (Figure 6C-F), but the details differed. First, the African elephant trunk modules had fewer and thicker myelin stripes (Figure 6D, F); this is a most interesting observation since African elephants have fewer trunk folds than Asian elephants. Second, African elephant trunk modules were significantly longer but not wider (Figure 6G). Asian elephant trunk modules had a much more belly-shaped appearance than African elephant trunk modules. The greatest width of the Asian elephant trunk module is at positions representing the trunk wrapping zone, which we determined from photographs. African elephant trunk modules were also significantly more ‘top-heavy’, i.e. they had their widest point much closer to the putative trunk tip (Figure 6H). We suggest that the shape differences between African and Asian elephant trunk modules might be related to the different grasping strategies of these two elephant species (Racine, 1980). African elephants have two fingers and tend to pinch objects (Figure 6I; upper), a grasping strategy that emphasizes the trunk tip in line with their ‘top-heavy’ trunk module. Asian elephants in contrast have only one finger and tend to wrap objects with their trunk (Racine, 1980; Figure 5I; lower). This grasping strategy engages more of the trunk and, in line with this behavior, the width of the Asian elephant trigeminal nucleus i is maximal in the trunk wrapping area.

Revised version: We found that the trigeminal bumps on the ventral brainstem differ significantly between African and Asian elephants (Figure 7A). We counted neurons in the principalis trunk module and found that African elephants (740210 ± 51902, mean ± SD) had more neurons than Asian elephants (636447 ± 69729, mean ± SD) and also had a larger volume principalis trunk module (Figure 7B). Supplementary Tables 2 and 3 provide further information on our counts of the trigeminal nuclei. As noted there, the dorsal finger accounted for a large fraction (∼20%) of the trunk modules. We wondered, why brainstem bumps differed between African and Asian elephants, and therefore closely investigated the shape of trunk modules in these species. A cytochrome-oxidase stained coronal section through the trunk module of the African elephant Indra is shown in Figure 7C. Drawings of coronal sections from this trunk module and the trunk module of other African elephants are shown in Figure 7D; myelin stripes (violet) were visible as whitish omissions of the cytochrome oxidase or of Nissl stains. We also determined the length, the width, and the longitudinal position of the greatest width of the module (black line; Figure 7D). A cytochrome-oxidase- stained coronal section through the trunk module of Asian elephant Raj is shown in Figure 7E and drawings from this and other trunk modules of Asian elephants are shown in Figure 7F. Basic aspects of the module trunk were similar in African and Asian elephants (Figure 7C-F), but the details differed. First, the African elephant trunk modules had fewer and thicker myelin stripes (Figure 7D, F); this is a most interesting observation since African elephants have fewer trunk folds than Asian elephants.

Differences between the putative trigeminal trunk modules of Asian (Elephas maximus) and African (Loxodonta africana) elephants

A, left, ventral view of the brainstem of African elephant Indra. Right, ventral view of the brainstem of Asian elephant Dumba. The anterior-posterior length of the trigeminal nuclei from the pons is indicated as a black line. The average trigeminal nuclei length refers to 10 African and 10 Asian trigeminal nuclei, the p-value refers to a Mann- Whitney test. The trigeminal nuclei bump is more elongated in African than in Asian elephants.

B, neuron number (left) and volume (right) of the putative principalis trunk module in Asian and African elephants. Trigeminal nuclei come from three African and three Asian elephants; p-values refer to unpaired t-tests.

C, micrograph of a cytochrome-oxidase-stained section through the putative trunk module of African elephant Indra.

D, drawings of the outlines and myelin stripes from cytochrome oxidase or Nissl stained sections through the putative trunk module of African elephants; the top drawing was made from the micrograph shown in C. The black line refers to the point of greatest width along the direction of myelin stripes on the putative trunk shaft (the putative finger was not considered in the width analysis).

E, micrograph of a cytochrome-oxidase-stained section through the putative trunk module of Asian elephant Raj.

F, drawings, of putative trunk modules from Asian elephants; conventions as in D.

G, length, and width of putative trunk module in African and Asian elephants. p-values refer to t-tests.

H, upper, drawing of the trunk of an African elephant. Lower, drawing of the trunk of an Asian elephant; note that Asian elephants have more folds. Arrows mark the point of greatest width of the putative trunk module (black lines in D, F) projected back on trunk positions in African (upper) and Asian (lower) elephant trunks. We also highlighted in color the dorsal (red) and ventral (pink) trunk tip and wrapping zone of Asian elephants in green and the analogous trunk part of African elephants in light green. The extent of the trunk wrapping zone was determined from photographs of Asian elephants wrapping objects. Specifically, we defined the wrapping zone as the trunk parts in contact with large objects (mangos, melons, fodder beets) being wrapped.

I, object grasping/pinching behavior in African (upper) and object wrapping strategy in Asian (lower) elephants (adapted from Kaufmann et al., 2022).

Added in the revised version: Asian elephants have more of these folds, but they are also more shallow and less pronounced (Schultz et al., 2023). Given our limited material and that myelin stripes are less conspicuous in Asian elephants than in African elephants, we were not able to ascertain the one-to-one match of stripes and folds that we could show for the African elephant Indra in Figure 5. Second, African elephant trunk modules were significantly longer but not wider (Figure 7G). Asian elephant trunk modules had a much more round-shaped appearance than African elephant trunk modules. The greatest width of the Asian elephant trunk module is at positions representing the trunk wrapping zone, which we determined from photographs. African elephant trunk modules were also significantly more ‘top-heavy’, i.e. they had their widest point much closer to the putative trunk tip (Figure 7H). We suggest that the shape differences between African and Asian elephant trunk modules might be related to the different grasping strategies of these two elephant species (Racine, 1980). African elephants have two fingers and tend to pinch objects (Figure 7I; upper), a grasping strategy that emphasizes the trunk tip in line with their ‘top-heavy’ trunk module. Asian elephants in contrast have only one finger and tend to wrap objects with their trunk (Racine, 1980; Figure 5I; lower). This grasping strategy engages more of the trunk and, in line with this behavior, the width of the Asian elephant trigeminal nucleus is maximal in the trunk wrapping area.

Discussion

Summary

We describe a pair of large bumps on the ventral surface of the elephant medulla that contain metabolically highly active, densely vascularized repeating modules. The trunk module contains an accurate myelin map of trunk folds. Mapping myelin stripes to the trunk folds indicated an increase in sensory magnification from the proximal to the distal trunk. Magnification analysis also identified an enlarged trunk wrapping zone in Asian elephants, who wrap objects with their trunk.

The ventral brainstem bumps likely correspond to elephant trigeminal nuclei

Earlier version: Establishing elephant brainstem organization is challenging because both tracing methods and in vivo electrophysiology cannot be applied to elephants. Our assignments of trigeminal nuclei deviate from earlier suggestions (Shoshani et al., 2006; Maseko et al., 2013), which assigned the putative trigeminal nuclei as inferior olive, and the structure identified as inferior olive, as trigeminal nuclei. For several reasons we think the ventral brainstem bumps correspond to the elephant trigeminal nuclei: 1. Their intense cytochrome oxidase reactivity, which is characteristic of the trigeminal nuclei of tactile specialists (Belford & Killackey, 1979; Ma, 1991; Catania bet al. 2013; Sawyer et al. 2014; Sawyer et al. 2015). 2. The large size of the putative trunk module, which matches with large infraorbital nerve of elephants (Purkart et al., 2022). 3. The elongation of the putative trunk module, which aligns with the trunk shape. 4. The arrangement of putative trigeminal modules (trunk module, nostril module, lower jaw module) that matches with elephant head topography. 5. The matching of module myelin stripes to trunk folds. 6. The position of putative trigeminal nuclei behind the pons, matches the position of the trigeminal nucleus in other mammals. 7. The organization of the trigeminal nuclei in modules repeating in the anterior-posterior direction, which is similar to other mammals (Belford & Killackey, 1979; Ma, 1991).

Revised version: Establishing elephant brainstem organization is challenging because both tracing methods and in vivo electrophysiology cannot be applied to elephants. Our assignments of trigeminal nuclei deviate from earlier suggestions (Shoshani et al., 2006; Maseko et al., 2013; Verhaart and Kramer, 1958; Verhaart 1962), which assigned the putative trigeminal nuclei as inferior olive, and the structure identified as inferior olive, as trigeminal nuclei. For several reasons we think the ventral brainstem bumps correspond to the elephant trigeminal nuclei: 1. Their intense cytochrome oxidase reactivity, which is characteristic of the trigeminal nuclei of tactile specialists (Belford & Killackey, 1979; Ma, 1991; Catania bet al. 2013; Sawyer et al. 2014; Sawyer et al. 2015). 2. The large size of the putative trunk module, matches with the large infraorbital nerve of elephants (Purkart et al., 2022). 3. The elongation of the putative trunk module, which aligns with the trunk shape. 4. The arrangement of putative trigeminal modules (trunk module, nostril module, lower jaw module) that matches with elephant head topography. 5. The matching of module myelin stripes to trunk folds. 6. The organization of the trigeminal nuclei in modules repeats in the anterior-posterior direction, similar to other mammals (Belford & Killackey, 1979; Ma, 1991).

Earlier version: The idea that ventral brainstem bumps correspond to inferior olive and that the laterally placed nuclei (referred to as inferior olive by us) correspond to the trigeminal nuclei (Shoshani et al., 2006; Maseko et al., 2013) is not viable: 1. The lateral nuclei (referred to as inferior olive by us) have the characteristic laminar cellular arrangement seen in the inferior olive of all mammals. 2. There are no indications of trigeminal-like modules in the lateral nuclei. 3. The lateral nuclei are not organized in modules repeating in the anterior-posterior direction, an arrangement seen in the trigeminal nuclei of other mammals 4. The ventral brainstem bumps (referred to as trigeminal nuclei by us) do not have the characteristic laminar cellular arrangement seen in the inferior olive of all mammals. 5. The very intense staining for cytochrome oxidase reactivity, seen in the ventral brainstem bumps, is not characteristic of the mammalian inferior olive. A synopsis of the above arguments suggests that the elephant trigeminal nuclei very likely correspond to the ventral brainstem bumps.

Revised version: The idea that ventral brainstem bumps correspond to inferior olive and that the laterally placed nuclei (referred to as inferior olive by us) correspond to the trigeminal nuclei (Shoshani et al., 2006; Maseko et al., 2013; Verhaart and Kramer, 1958; Verhaart 1962) appears unlikely: 1. The lateral nuclei (referred to as inferior olive by us) have the characteristic laminar cellular arrangement seen in the inferior olive of all mammals. 2. There are no indications of trigeminal-like modules in the lateral nuclei. 3. The lateral nuclei are not organized in modules repeating in the anterior-posterior direction, an arrangement seen in the trigeminal nuclei of other mammals 4. The ventral brainstem bumps (referred to as trigeminal nuclei by us) do not have the characteristic laminar cellular arrangement seen in the inferior olive of all mammals. 5. The very intense staining for cytochrome oxidase reactivity, seen in the ventral brainstem bumps, is not characteristic of the mammalian inferior olive. 6. The absence of climbing fiber bundles in the ventral brainstem bumps (Figure 2), which lets it appear unlikely that these bumps correspond to the inferior olive.

Added in the revised version: A synopsis of the above arguments suggests that the elephant trigeminal nuclei very likely correspond to the ventral brainstem bumps (please also see our Referee Figure 3 and Referee Table 1 in our response letter published along with the article). As clearly expressed by the comments of Referee 2 published along with our article, the strength other partitioning scheme, which assign the ventral brainstem bump as inferior olive is the better positional match of the assignment with the inferior olive of other mammals. We think this is a valid criticism. Our model offers single anatomical nuclei with trigeminal nuclei that look like trigeminal nuclei of other mammals, a trunk module with a striking resemblance to the trunk, and an inferior olive that looks like the inferior olive of other mammals.

Detailed neuroanatomic mapping of elephant trigeminal nuclei

Our work provided detailed mapping of the elephant trigeminal brainstem into four repeating nuclei, consisting of several facial modules (most prominently the trunk module, the nostril module, and the lower jaw module). Because of the level of detail of our topography suggestions, however, we think it will be relatively straightforward to test the validity of our topography suggestions. Specifically, we predict that the dorsal trunk finger representation of the principalis trunk module will be connected with the distal part of the dorsal subnucleus of the facial nucleus (which contains the putative motor representation of the dorsal trunk finger (Kaufmann et al., 2022)). We would also predict that the dorsal trunk finger representation of the principalis trunk module will be connected with the dorsal trunk finger representation of the oralis nucleus. These connectivity suggestions are in the 1-2 cm range and can be tested with postmortem tracers like DiI. We also predict fewer myelin stripes in trunk modules of elephants with particularly few trunk folds (newborn or fetuses of African elephants) compared to adult African elephants.

A myelin map of trunk folds, white matter function, and myeloarchitecture

According to conventional wisdom, neurons (gray matter) are the site of processing, and myelinated axons (white matter) are a subjugated supply system, whose sole function is to bring the correct axonal input to neurons. Our observations on trunk module myelin stripes are at odds with this view of myelin. Specifically, myelin stripes show no tapering (which we would expect if axons divert off into the tissue). More than that, there is no correlation between myelin stripe thickness (which presumably correlates with axon numbers) and trigeminal module neuron numbers. Thus, there are numerous myelinated axons, where we observe few or no trigeminal neurons. These observations are incompatible with the idea that myelin stripes form an axonal ‘supply’ system or that their prime function is to connect neurons. What do myelin stripe axons do, if they do not connect neurons? We suggest that myelin stripes serve to separate rather than connect neurons. Specifically, trunk module myelin stripes look like a map of trunk folds. Myelin stripes match with the number, orientation, and species-specific patterning of trunk folds. We note that if myelin stripes would behave as an axonal ‘supply’ system, they would be very thin/invisible in the proximal trunk, proximal trunk folds would not be visible, and distal stripes should be very thick. If myelinated axons have a life of their own and do not simply go where they find target neurons, we need to analyze them in novel ways. In particular, it seems to be a good idea to ‘look’ at patterns of myelination, rather than to immediately assume that this is a supply/connectivity system. We note early neuroanatomists like Oskar Vogt (Vogt, 1911; Niewenhuys, Broere & Cerliani, 2015) described incredibly intricate patterns of intracortical myeloarchitecture, patterns that are not easily explained in terms of a connectivity system to this day. In conclusion, we propose a novel white-matter function, which is to separate and functionally demarcate neurons as opposed to the conventionally assumed white-matter function of connecting neurons.

Trigeminal organization in Asian and African elephants

At first sight, Asian and African elephant trigeminal nuclei are very similar. Both elephants have big ventral brainstem bumps, which contain the same modules (a putative trunk, nostril, and lower jaw module) and the nuclei stain intensely for cytochrome oxidase reactivity. A closer look reveals species differences, however, which may relate to the different trunk grasping strategies of Asian and African elephants. The first difference refers to the shape of the ventral brainstem bump, which is more roundish and shorter in Asian elephants and more elongated in African elephants. To our knowledge, this ventral brainstem bump is the only hitherto described difference, which allows us to differentiate Asian and African elephant brains from the outside. The length difference between Asian and African elephant ventral bumps reflects the different shapes of Asian and African elephant trunk modules. The African elephant trunk module is notably longer, more slender, and more top-heavy than the swaged Asian elephant trunk module. As we pointed out in Figure 5, such differences imply an enlargement of the trunk wrapping zone in the Asian elephant trunk module, in line with the object-wrapping behavior of Asian elephants (Racine 1980). Similarly, the top-heavy shape of the African elephant trunk module could be instrumental in the object-pinching of African elephants (Racine, 1980). These trigeminal differences are reminiscent of similar Asian-African species differences in the elephant facial nucleus (Kaufmann et al., 2022). We conclude that grasping behavior shapes the species-specific architecture of the trigeminal nuclei.

Conclusion

The elephant brainstem is exquisitely well-ordered and contains very large and detailed trigeminal representations. Trunk module myelin stripes form a map of trunk folds and accordingly serve to functionally separate neurons rather than to connect them. Further work should test the predictions of sensory topographies outlined here and ask, what further insights the elephant brain provides about the organization of gray and white matter.